./vampire_rel_mtpa-gnn_8867 -t 300 -i 100000 --input_syntax tptp -sa otter -ncem /home/sudamar2/mtpa-gnn/tptpOverfit100k3/loop8/script-model.pt -npcc on -si on -rtra on --random_seed 9 Problems/CAT/CAT027+3.p -p on
% Refutation found. Thanks to Tanya!
% SZS status Theorem for CAT027+3
% SZS output start Proof for CAT027+3
1394. ! [X2,X1,X0] : (m2_relset_1(X2,X0,X1) <=> m1_relset_1(X2,X0,X1)) [input(axiom)]
6631. ! [X0] : (l1_cat_1(X0) => ~v1_xboole_0(u1_cat_1(X0))) [input(axiom)]
6632. ! [X0] : (l1_cat_1(X0) => ~v1_xboole_0(u2_cat_1(X0))) [input(axiom)]
8397. ! [X1,X0] : ((v2_cat_1(X0) & l1_cat_1(X1) & l1_cat_1(X0) & v2_cat_1(X1)) => (v2_cat_1(k11_cat_2(X0,X1)) & l1_cat_1(k11_cat_2(X0,X1)))) [input(axiom)]
8405. ! [X1,X0] : ((v2_cat_1(X0) & l1_cat_1(X0) & v2_cat_1(X1) & l1_cat_1(X1)) => m2_cat_1(k17_cat_2(X0,X1),k11_cat_2(X0,X1),X1)) [input(axiom)]
10520. ! [X1,X0,X2,X3] : ((l1_cat_1(X0) & m2_cat_1(X2,X0,X1) & l1_cat_1(X1) & m2_cat_1(X3,X0,X1) & v2_cat_1(X0) & v2_cat_1(X1)) => ! [X4] : (m1_nattra_1(X4,X0,X1,X2,X3) => (m2_relset_1(X4,u1_cat_1(X0),u2_cat_1(X1)) & v1_funct_1(X4) & v1_funct_2(X4,u1_cat_1(X0),u2_cat_1(X1))))) [input(axiom)]
10522. ! [X3,X0,X1,X2] : ((m2_cat_1(X2,X0,X1) & m2_cat_1(X3,X0,X1) & l1_cat_1(X1) & v2_cat_1(X0) & v2_cat_1(X1) & l1_cat_1(X0)) => ! [X4] : (m2_nattra_1(X4,X0,X1,X2,X3) => m1_nattra_1(X4,X0,X1,X2,X3))) [input(axiom)]
10531. ! [X3,X2,X4,X1,X0,X5] : ((v1_funct_2(X4,X0,X1) & ~v1_xboole_0(X3) & v1_funct_2(X5,X2,X3) & ~v1_xboole_0(X0) & v1_funct_1(X4) & ~v1_xboole_0(X2) & ~v1_xboole_0(X1) & m1_relset_1(X5,X2,X3) & v1_funct_1(X5) & m1_relset_1(X4,X0,X1)) => (r4_nattra_1(X0,X1,X2,X3,X4,X5) => r4_nattra_1(X0,X1,X2,X3,X5,X4))) [input(axiom)]
10542. ! [X0,X2,X1] : ((l1_cat_1(X0) & v2_cat_1(X0) & v2_cat_1(X1) & m2_cat_1(X2,X0,X1) & l1_cat_1(X1)) => m2_nattra_1(k7_nattra_1(X0,X1,X2),X0,X1,X2,X2)) [input(axiom)]
10543. ! [X0,X2,X1] : ((v2_cat_1(X0) & l1_cat_1(X1) & v2_cat_1(X1) & l1_cat_1(X0) & m2_cat_1(X2,X0,X1)) => k4_nattra_1(X0,X1,X2) = k7_nattra_1(X0,X1,X2)) [input(axiom)]
11397. ! [X0] : ((v2_cat_1(X0) & l1_cat_1(X0)) => ! [X1] : ((v2_cat_1(X1) & l1_cat_1(X1)) => ! [X2] : ((l1_cat_1(X2) & v2_cat_1(X2)) => ! [X3] : (m2_cat_1(X3,X0,X1) => ! [X4] : (m2_cat_1(X4,X1,X2) => r4_nattra_1(u1_cat_1(X0),u2_cat_1(X2),u1_cat_1(X0),u2_cat_1(X2),k6_isocat_1(X0,X1,X2,X3,X3,k7_nattra_1(X0,X1,X3),X4),k7_nattra_1(X0,X2,k2_isocat_1(X0,X1,X2,X3,X4)))))))) [input(axiom)]
11434. ! [X1,X4,X3,X5,X2,X6,X0] : ((v2_cat_1(X0) & l1_cat_1(X0) & l1_cat_1(X1) & m2_nattra_1(X5,X0,X1,X3,X4) & l1_cat_1(X2) & m2_cat_1(X4,X0,X1) & v2_cat_1(X1) & m2_cat_1(X3,X0,X1) & v2_cat_1(X2) & m2_cat_1(X6,X1,X2)) => m2_nattra_1(k6_isocat_1(X0,X1,X2,X3,X4,X5,X6),X0,X2,k2_isocat_1(X0,X1,X2,X3,X6),k2_isocat_1(X0,X1,X2,X4,X6))) [input(axiom)]
11735. ! [X1,X0] : ((v2_cat_1(X1) & l1_cat_1(X0) & l1_cat_1(X1) & v2_cat_1(X0)) => m2_cat_1(k8_isocat_2(X0,X1),k11_cat_2(X0,X1),X0)) [input(axiom)]
11736. ! [X1,X0] : ((v2_cat_1(X1) & l1_cat_1(X0) & l1_cat_1(X1) & v2_cat_1(X0)) => k16_cat_2(X0,X1) = k8_isocat_2(X0,X1)) [input(axiom)]
11738. ! [X1,X0] : ((v2_cat_1(X1) & l1_cat_1(X1) & v2_cat_1(X0) & l1_cat_1(X0)) => k17_cat_2(X0,X1) = k9_isocat_2(X0,X1)) [input(axiom)]
11741. ! [X2,X1,X3,X0] : ((v2_cat_1(X2) & l1_cat_1(X0) & l1_cat_1(X1) & v2_cat_1(X1) & v2_cat_1(X0) & m2_cat_1(X3,X0,k11_cat_2(X1,X2)) & l1_cat_1(X2)) => m2_cat_1(k11_isocat_2(X0,X1,X2,X3),X0,X1)) [input(axiom)]
11742. ! [X0,X3,X1,X2] : ((l1_cat_1(X0) & v2_cat_1(X2) & m2_cat_1(X3,X0,k11_cat_2(X1,X2)) & v2_cat_1(X0) & l1_cat_1(X1) & l1_cat_1(X2) & v2_cat_1(X1)) => m2_cat_1(k12_isocat_2(X0,X1,X2,X3),X0,X2)) [input(axiom)]
11743. ! [X2,X4,X5,X1,X3,X0] : ((m2_cat_1(X3,X0,k11_cat_2(X1,X2)) & l1_cat_1(X2) & l1_cat_1(X1) & m2_nattra_1(X5,X0,k11_cat_2(X1,X2),X3,X4) & v2_cat_1(X1) & v2_cat_1(X0) & m2_cat_1(X4,X0,k11_cat_2(X1,X2)) & l1_cat_1(X0) & v2_cat_1(X2)) => m2_nattra_1(k13_isocat_2(X0,X1,X2,X3,X4,X5),X0,X1,k11_isocat_2(X0,X1,X2,X3),k11_isocat_2(X0,X1,X2,X4))) [input(axiom)]
11744. ! [X3,X1,X2,X5,X4,X0] : ((m2_cat_1(X4,X0,k11_cat_2(X1,X2)) & v2_cat_1(X1) & m2_cat_1(X3,X0,k11_cat_2(X1,X2)) & v2_cat_1(X0) & v2_cat_1(X2) & l1_cat_1(X1) & l1_cat_1(X0) & m2_nattra_1(X5,X0,k11_cat_2(X1,X2),X3,X4) & l1_cat_1(X2)) => m2_nattra_1(k14_isocat_2(X0,X1,X2,X3,X4,X5),X0,X2,k12_isocat_2(X0,X1,X2,X3),k12_isocat_2(X0,X1,X2,X4))) [input(axiom)]
11791. ! [X0] : ((v2_cat_1(X0) & l1_cat_1(X0)) => ! [X1] : ((v2_cat_1(X1) & l1_cat_1(X1)) => ! [X2] : ((v2_cat_1(X2) & l1_cat_1(X2)) => ! [X3] : (m2_cat_1(X3,X0,k11_cat_2(X1,X2)) => k11_isocat_2(X0,X1,X2,X3) = k2_isocat_1(X0,k11_cat_2(X1,X2),X1,X3,k8_isocat_2(X1,X2)))))) [input(axiom)]
11792. ! [X0] : ((v2_cat_1(X0) & l1_cat_1(X0)) => ! [X1] : ((v2_cat_1(X1) & l1_cat_1(X1)) => ! [X2] : ((v2_cat_1(X2) & l1_cat_1(X2)) => ! [X3] : (m2_cat_1(X3,X0,k11_cat_2(X1,X2)) => k12_isocat_2(X0,X1,X2,X3) = k2_isocat_1(X0,k11_cat_2(X1,X2),X2,X3,k9_isocat_2(X1,X2)))))) [input(axiom)]
11795. ! [X0] : ((l1_cat_1(X0) & v2_cat_1(X0)) => ! [X1] : ((l1_cat_1(X1) & v2_cat_1(X1)) => ! [X2] : ((l1_cat_1(X2) & v2_cat_1(X2)) => ! [X3] : (m2_cat_1(X3,X0,k11_cat_2(X1,X2)) => ! [X4] : (m2_cat_1(X4,X0,k11_cat_2(X1,X2)) => ! [X5] : (m2_nattra_1(X5,X0,k11_cat_2(X1,X2),X3,X4) => k13_isocat_2(X0,X1,X2,X3,X4,X5) = k6_isocat_1(X0,k11_cat_2(X1,X2),X1,X3,X4,X5,k8_isocat_2(X1,X2)))))))) [input(axiom)]
11796. ! [X0] : ((l1_cat_1(X0) & v2_cat_1(X0)) => ! [X1] : ((v2_cat_1(X1) & l1_cat_1(X1)) => ! [X2] : ((v2_cat_1(X2) & l1_cat_1(X2)) => ! [X3] : (m2_cat_1(X3,X0,k11_cat_2(X1,X2)) => ! [X4] : (m2_cat_1(X4,X0,k11_cat_2(X1,X2)) => ! [X5] : (m2_nattra_1(X5,X0,k11_cat_2(X1,X2),X3,X4) => k14_isocat_2(X0,X1,X2,X3,X4,X5) = k6_isocat_1(X0,k11_cat_2(X1,X2),X2,X3,X4,X5,k9_isocat_2(X1,X2)))))))) [input(axiom)]
11799. ! [X0] : ((l1_cat_1(X0) & v2_cat_1(X0)) => ! [X1] : ((l1_cat_1(X1) & v2_cat_1(X1)) => ! [X2] : ((l1_cat_1(X2) & v2_cat_1(X2)) => ! [X3] : (m2_cat_1(X3,X0,k11_cat_2(X1,X2)) => (r4_nattra_1(u1_cat_1(X0),u2_cat_1(X1),u1_cat_1(X0),u2_cat_1(X1),k7_nattra_1(X0,X1,k11_isocat_2(X0,X1,X2,X3)),k13_isocat_2(X0,X1,X2,X3,X3,k7_nattra_1(X0,k11_cat_2(X1,X2),X3))) & r4_nattra_1(u1_cat_1(X0),u2_cat_1(X2),u1_cat_1(X0),u2_cat_1(X2),k7_nattra_1(X0,X2,k12_isocat_2(X0,X1,X2,X3)),k14_isocat_2(X0,X1,X2,X3,X3,k7_nattra_1(X0,k11_cat_2(X1,X2),X3)))))))) [input(conjecture)]
11800. ~! [X0] : ((l1_cat_1(X0) & v2_cat_1(X0)) => ! [X1] : ((l1_cat_1(X1) & v2_cat_1(X1)) => ! [X2] : ((l1_cat_1(X2) & v2_cat_1(X2)) => ! [X3] : (m2_cat_1(X3,X0,k11_cat_2(X1,X2)) => (r4_nattra_1(u1_cat_1(X0),u2_cat_1(X1),u1_cat_1(X0),u2_cat_1(X1),k7_nattra_1(X0,X1,k11_isocat_2(X0,X1,X2,X3)),k13_isocat_2(X0,X1,X2,X3,X3,k7_nattra_1(X0,k11_cat_2(X1,X2),X3))) & r4_nattra_1(u1_cat_1(X0),u2_cat_1(X2),u1_cat_1(X0),u2_cat_1(X2),k7_nattra_1(X0,X2,k12_isocat_2(X0,X1,X2,X3)),k14_isocat_2(X0,X1,X2,X3,X3,k7_nattra_1(X0,k11_cat_2(X1,X2),X3)))))))) [negated conjecture 11799]
11820. ! [X0,X1,X4,X3,X2,X5] : ((m2_nattra_1(X2,X5,k11_cat_2(X3,X0),X4,X1) & l1_cat_1(X3) & l1_cat_1(X5) & l1_cat_1(X0) & v2_cat_1(X3) & v2_cat_1(X0) & m2_cat_1(X1,X5,k11_cat_2(X3,X0)) & m2_cat_1(X4,X5,k11_cat_2(X3,X0)) & v2_cat_1(X5)) => m2_nattra_1(k13_isocat_2(X5,X3,X0,X4,X1,X2),X5,X3,k11_isocat_2(X5,X3,X0,X4),k11_isocat_2(X5,X3,X0,X1))) [rectify 11743]
12033. ! [X1,X0] : ((l1_cat_1(X0) & l1_cat_1(X1) & v2_cat_1(X1) & v2_cat_1(X0)) => k8_isocat_2(X1,X0) = k16_cat_2(X1,X0)) [rectify 11736]
12800. ! [X1,X2,X0] : (m1_relset_1(X0,X2,X1) <=> m2_relset_1(X0,X2,X1)) [rectify 1394]
12827. ! [X0,X1,X2,X3] : ((v2_cat_1(X2) & v2_cat_1(X3) & m2_cat_1(X1,X0,k11_cat_2(X2,X3)) & v2_cat_1(X0) & l1_cat_1(X0) & l1_cat_1(X3) & l1_cat_1(X2)) => m2_cat_1(k12_isocat_2(X0,X2,X3,X1),X0,X3)) [rectify 11742]
12844. ! [X1,X4,X0,X2,X5,X3] : ((m1_relset_1(X2,X4,X3) & v1_funct_1(X5) & ~v1_xboole_0(X3) & ~v1_xboole_0(X0) & m1_relset_1(X5,X1,X0) & v1_funct_2(X2,X4,X3) & ~v1_xboole_0(X1) & v1_funct_1(X2) & ~v1_xboole_0(X4) & v1_funct_2(X5,X1,X0)) => (r4_nattra_1(X4,X3,X1,X0,X2,X5) => r4_nattra_1(X4,X3,X1,X0,X5,X2))) [rectify 10531]
13347. ! [X1,X6,X3,X0,X5,X2,X4] : ((l1_cat_1(X6) & m2_nattra_1(X3,X6,X0,X2,X1) & v2_cat_1(X4) & v2_cat_1(X6) & l1_cat_1(X0) & m2_cat_1(X5,X0,X4) & v2_cat_1(X0) & l1_cat_1(X4) & m2_cat_1(X2,X6,X0) & m2_cat_1(X1,X6,X0)) => m2_nattra_1(k6_isocat_1(X6,X0,X4,X2,X1,X3,X5),X6,X4,k2_isocat_1(X6,X0,X4,X2,X5),k2_isocat_1(X6,X0,X4,X1,X5))) [rectify 11434]
13402. ! [X0,X1] : ((l1_cat_1(X1) & l1_cat_1(X0) & v2_cat_1(X0) & v2_cat_1(X1)) => k17_cat_2(X1,X0) = k9_isocat_2(X1,X0)) [rectify 11738]
13411. ! [X3,X0,X1,X2] : ((v2_cat_1(X1) & l1_cat_1(X1) & m2_cat_1(X3,X1,X0) & v2_cat_1(X0) & l1_cat_1(X0) & m2_cat_1(X2,X1,X0)) => ! [X4] : (m1_nattra_1(X4,X1,X0,X2,X3) => (m2_relset_1(X4,u1_cat_1(X1),u2_cat_1(X0)) & v1_funct_1(X4) & v1_funct_2(X4,u1_cat_1(X1),u2_cat_1(X0))))) [rectify 10520]
13499. ! [X1,X0] : ((v2_cat_1(X1) & l1_cat_1(X1) & v2_cat_1(X0) & l1_cat_1(X0)) => m2_cat_1(k17_cat_2(X1,X0),k11_cat_2(X1,X0),X0)) [rectify 8405]
13567. ! [X1,X0,X3,X2] : ((m2_cat_1(X3,X1,X2) & l1_cat_1(X2) & m2_cat_1(X0,X1,X2) & v2_cat_1(X2) & v2_cat_1(X1) & l1_cat_1(X1)) => ! [X4] : (m2_nattra_1(X4,X1,X2,X3,X0) => m1_nattra_1(X4,X1,X2,X3,X0))) [rectify 10522]
13609. ! [X1,X0,X5,X4,X3,X2] : ((m2_cat_1(X4,X5,k11_cat_2(X1,X2)) & v2_cat_1(X2) & l1_cat_1(X1) & l1_cat_1(X5) & m2_cat_1(X0,X5,k11_cat_2(X1,X2)) & v2_cat_1(X5) & v2_cat_1(X1) & l1_cat_1(X2) & m2_nattra_1(X3,X5,k11_cat_2(X1,X2),X0,X4)) => m2_nattra_1(k14_isocat_2(X5,X1,X2,X0,X4,X3),X5,X2,k12_isocat_2(X5,X1,X2,X0),k12_isocat_2(X5,X1,X2,X4))) [rectify 11744]
14132. ! [X1,X2,X0] : ((v2_cat_1(X2) & l1_cat_1(X0) & v2_cat_1(X0) & l1_cat_1(X2) & m2_cat_1(X1,X0,X2)) => k7_nattra_1(X0,X2,X1) = k4_nattra_1(X0,X2,X1)) [rectify 10543]
14159. ! [X1,X3,X0,X2] : ((l1_cat_1(X1) & v2_cat_1(X0) & l1_cat_1(X3) & m2_cat_1(X2,X3,k11_cat_2(X1,X0)) & l1_cat_1(X0) & v2_cat_1(X1) & v2_cat_1(X3)) => m2_cat_1(k11_isocat_2(X3,X1,X0,X2),X3,X1)) [rectify 11741]
14421. ! [X1,X2,X0] : ((m2_cat_1(X1,X0,X2) & l1_cat_1(X2) & l1_cat_1(X0) & v2_cat_1(X0) & v2_cat_1(X2)) => m2_nattra_1(k7_nattra_1(X0,X2,X1),X0,X2,X1,X1)) [rectify 10542]
14876. ! [X1,X0] : (m2_cat_1(k8_isocat_2(X0,X1),k11_cat_2(X0,X1),X0) | (~v2_cat_1(X1) | ~l1_cat_1(X0) | ~l1_cat_1(X1) | ~v2_cat_1(X0))) [ennf transformation 11735]
14877. ! [X0,X1] : (~l1_cat_1(X1) | ~v2_cat_1(X0) | m2_cat_1(k8_isocat_2(X0,X1),k11_cat_2(X0,X1),X0) | ~v2_cat_1(X1) | ~l1_cat_1(X0)) [flattening 14876]
15340. ! [X0,X1] : (k17_cat_2(X1,X0) = k9_isocat_2(X1,X0) | (~l1_cat_1(X1) | ~l1_cat_1(X0) | ~v2_cat_1(X0) | ~v2_cat_1(X1))) [ennf transformation 13402]
15341. ! [X0,X1] : (~v2_cat_1(X1) | ~v2_cat_1(X0) | k17_cat_2(X1,X0) = k9_isocat_2(X1,X0) | ~l1_cat_1(X1) | ~l1_cat_1(X0)) [flattening 15340]
15396. ! [X1,X0,X3,X2] : (! [X4] : (m1_nattra_1(X4,X1,X2,X3,X0) | ~m2_nattra_1(X4,X1,X2,X3,X0)) | (~m2_cat_1(X3,X1,X2) | ~l1_cat_1(X2) | ~m2_cat_1(X0,X1,X2) | ~v2_cat_1(X2) | ~v2_cat_1(X1) | ~l1_cat_1(X1))) [ennf transformation 13567]
15397. ! [X1,X3,X2,X0] : (~l1_cat_1(X1) | ~m2_cat_1(X0,X1,X2) | ~v2_cat_1(X2) | ~m2_cat_1(X3,X1,X2) | ~v2_cat_1(X1) | ~l1_cat_1(X2) | ! [X4] : (m1_nattra_1(X4,X1,X2,X3,X0) | ~m2_nattra_1(X4,X1,X2,X3,X0))) [flattening 15396]
16789. ! [X1,X0] : ((v2_cat_1(k11_cat_2(X0,X1)) & l1_cat_1(k11_cat_2(X0,X1))) | (~v2_cat_1(X0) | ~l1_cat_1(X1) | ~l1_cat_1(X0) | ~v2_cat_1(X1))) [ennf transformation 8397]
16790. ! [X1,X0] : (~l1_cat_1(X0) | ~v2_cat_1(X0) | ~l1_cat_1(X1) | (v2_cat_1(k11_cat_2(X0,X1)) & l1_cat_1(k11_cat_2(X0,X1))) | ~v2_cat_1(X1)) [flattening 16789]
16933. ! [X3,X0,X1,X2] : (! [X4] : (~m1_nattra_1(X4,X1,X0,X2,X3) | (m2_relset_1(X4,u1_cat_1(X1),u2_cat_1(X0)) & v1_funct_1(X4) & v1_funct_2(X4,u1_cat_1(X1),u2_cat_1(X0)))) | (~v2_cat_1(X1) | ~l1_cat_1(X1) | ~m2_cat_1(X3,X1,X0) | ~v2_cat_1(X0) | ~l1_cat_1(X0) | ~m2_cat_1(X2,X1,X0))) [ennf transformation 13411]
16934. ! [X0,X2,X1,X3] : (~m2_cat_1(X2,X1,X0) | ! [X4] : (~m1_nattra_1(X4,X1,X0,X2,X3) | (m2_relset_1(X4,u1_cat_1(X1),u2_cat_1(X0)) & v1_funct_1(X4) & v1_funct_2(X4,u1_cat_1(X1),u2_cat_1(X0)))) | ~l1_cat_1(X0) | ~l1_cat_1(X1) | ~v2_cat_1(X0) | ~m2_cat_1(X3,X1,X0) | ~v2_cat_1(X1)) [flattening 16933]
17305. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (~m2_cat_1(X4,X0,k11_cat_2(X1,X2)) | ! [X5] : (k13_isocat_2(X0,X1,X2,X3,X4,X5) = k6_isocat_1(X0,k11_cat_2(X1,X2),X1,X3,X4,X5,k8_isocat_2(X1,X2)) | ~m2_nattra_1(X5,X0,k11_cat_2(X1,X2),X3,X4))) | ~m2_cat_1(X3,X0,k11_cat_2(X1,X2))) | (~l1_cat_1(X2) | ~v2_cat_1(X2))) | (~l1_cat_1(X1) | ~v2_cat_1(X1))) | (~l1_cat_1(X0) | ~v2_cat_1(X0))) [ennf transformation 11795]
17306. ! [X0] : (~l1_cat_1(X0) | ! [X1] : (! [X2] : (! [X3] : (! [X4] : (~m2_cat_1(X4,X0,k11_cat_2(X1,X2)) | ! [X5] : (k13_isocat_2(X0,X1,X2,X3,X4,X5) = k6_isocat_1(X0,k11_cat_2(X1,X2),X1,X3,X4,X5,k8_isocat_2(X1,X2)) | ~m2_nattra_1(X5,X0,k11_cat_2(X1,X2),X3,X4))) | ~m2_cat_1(X3,X0,k11_cat_2(X1,X2))) | ~l1_cat_1(X2) | ~v2_cat_1(X2)) | ~l1_cat_1(X1) | ~v2_cat_1(X1)) | ~v2_cat_1(X0)) [flattening 17305]
17429. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (~m2_cat_1(X4,X0,k11_cat_2(X1,X2)) | ! [X5] : (~m2_nattra_1(X5,X0,k11_cat_2(X1,X2),X3,X4) | k14_isocat_2(X0,X1,X2,X3,X4,X5) = k6_isocat_1(X0,k11_cat_2(X1,X2),X2,X3,X4,X5,k9_isocat_2(X1,X2)))) | ~m2_cat_1(X3,X0,k11_cat_2(X1,X2))) | (~v2_cat_1(X2) | ~l1_cat_1(X2))) | (~v2_cat_1(X1) | ~l1_cat_1(X1))) | (~l1_cat_1(X0) | ~v2_cat_1(X0))) [ennf transformation 11796]
17430. ! [X0] : (~v2_cat_1(X0) | ~l1_cat_1(X0) | ! [X1] : (~v2_cat_1(X1) | ~l1_cat_1(X1) | ! [X2] : (! [X3] : (! [X4] : (~m2_cat_1(X4,X0,k11_cat_2(X1,X2)) | ! [X5] : (~m2_nattra_1(X5,X0,k11_cat_2(X1,X2),X3,X4) | k14_isocat_2(X0,X1,X2,X3,X4,X5) = k6_isocat_1(X0,k11_cat_2(X1,X2),X2,X3,X4,X5,k9_isocat_2(X1,X2)))) | ~m2_cat_1(X3,X0,k11_cat_2(X1,X2))) | ~l1_cat_1(X2) | ~v2_cat_1(X2)))) [flattening 17429]
19659. ! [X1,X0] : (m2_cat_1(k17_cat_2(X1,X0),k11_cat_2(X1,X0),X0) | (~v2_cat_1(X1) | ~l1_cat_1(X1) | ~v2_cat_1(X0) | ~l1_cat_1(X0))) [ennf transformation 13499]
19660. ! [X1,X0] : (~l1_cat_1(X1) | ~l1_cat_1(X0) | ~v2_cat_1(X1) | m2_cat_1(k17_cat_2(X1,X0),k11_cat_2(X1,X0),X0) | ~v2_cat_1(X0)) [flattening 19659]
20644. ! [X1,X0,X5,X4,X3,X2] : (m2_nattra_1(k14_isocat_2(X5,X1,X2,X0,X4,X3),X5,X2,k12_isocat_2(X5,X1,X2,X0),k12_isocat_2(X5,X1,X2,X4)) | (~m2_cat_1(X4,X5,k11_cat_2(X1,X2)) | ~v2_cat_1(X2) | ~l1_cat_1(X1) | ~l1_cat_1(X5) | ~m2_cat_1(X0,X5,k11_cat_2(X1,X2)) | ~v2_cat_1(X5) | ~v2_cat_1(X1) | ~l1_cat_1(X2) | ~m2_nattra_1(X3,X5,k11_cat_2(X1,X2),X0,X4))) [ennf transformation 13609]
20645. ! [X2,X0,X5,X1,X4,X3] : (~l1_cat_1(X2) | m2_nattra_1(k14_isocat_2(X5,X1,X2,X0,X4,X3),X5,X2,k12_isocat_2(X5,X1,X2,X0),k12_isocat_2(X5,X1,X2,X4)) | ~l1_cat_1(X1) | ~m2_cat_1(X0,X5,k11_cat_2(X1,X2)) | ~m2_cat_1(X4,X5,k11_cat_2(X1,X2)) | ~l1_cat_1(X5) | ~v2_cat_1(X5) | ~v2_cat_1(X2) | ~v2_cat_1(X1) | ~m2_nattra_1(X3,X5,k11_cat_2(X1,X2),X0,X4)) [flattening 20644]
21275. ! [X1,X0] : (k8_isocat_2(X1,X0) = k16_cat_2(X1,X0) | (~l1_cat_1(X0) | ~l1_cat_1(X1) | ~v2_cat_1(X1) | ~v2_cat_1(X0))) [ennf transformation 12033]
21276. ! [X0,X1] : (~v2_cat_1(X1) | ~v2_cat_1(X0) | ~l1_cat_1(X1) | ~l1_cat_1(X0) | k8_isocat_2(X1,X0) = k16_cat_2(X1,X0)) [flattening 21275]
22274. ! [X1,X2,X0] : (k7_nattra_1(X0,X2,X1) = k4_nattra_1(X0,X2,X1) | (~v2_cat_1(X2) | ~l1_cat_1(X0) | ~v2_cat_1(X0) | ~l1_cat_1(X2) | ~m2_cat_1(X1,X0,X2))) [ennf transformation 14132]
22275. ! [X0,X1,X2] : (k7_nattra_1(X0,X2,X1) = k4_nattra_1(X0,X2,X1) | ~v2_cat_1(X2) | ~l1_cat_1(X2) | ~v2_cat_1(X0) | ~m2_cat_1(X1,X0,X2) | ~l1_cat_1(X0)) [flattening 22274]
22294. ! [X0,X1,X4,X3,X2,X5] : (m2_nattra_1(k13_isocat_2(X5,X3,X0,X4,X1,X2),X5,X3,k11_isocat_2(X5,X3,X0,X4),k11_isocat_2(X5,X3,X0,X1)) | (~m2_nattra_1(X2,X5,k11_cat_2(X3,X0),X4,X1) | ~l1_cat_1(X3) | ~l1_cat_1(X5) | ~l1_cat_1(X0) | ~v2_cat_1(X3) | ~v2_cat_1(X0) | ~m2_cat_1(X1,X5,k11_cat_2(X3,X0)) | ~m2_cat_1(X4,X5,k11_cat_2(X3,X0)) | ~v2_cat_1(X5))) [ennf transformation 11820]
22295. ! [X5,X3,X0,X1,X4,X2] : (~m2_nattra_1(X2,X5,k11_cat_2(X3,X0),X4,X1) | ~m2_cat_1(X4,X5,k11_cat_2(X3,X0)) | ~v2_cat_1(X0) | m2_nattra_1(k13_isocat_2(X5,X3,X0,X4,X1,X2),X5,X3,k11_isocat_2(X5,X3,X0,X4),k11_isocat_2(X5,X3,X0,X1)) | ~l1_cat_1(X3) | ~m2_cat_1(X1,X5,k11_cat_2(X3,X0)) | ~v2_cat_1(X3) | ~l1_cat_1(X5) | ~l1_cat_1(X0) | ~v2_cat_1(X5)) [flattening 22294]
22625. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (! [X4] : (r4_nattra_1(u1_cat_1(X0),u2_cat_1(X2),u1_cat_1(X0),u2_cat_1(X2),k6_isocat_1(X0,X1,X2,X3,X3,k7_nattra_1(X0,X1,X3),X4),k7_nattra_1(X0,X2,k2_isocat_1(X0,X1,X2,X3,X4))) | ~m2_cat_1(X4,X1,X2)) | ~m2_cat_1(X3,X0,X1)) | (~l1_cat_1(X2) | ~v2_cat_1(X2))) | (~v2_cat_1(X1) | ~l1_cat_1(X1))) | (~v2_cat_1(X0) | ~l1_cat_1(X0))) [ennf transformation 11397]
22626. ! [X0] : (~l1_cat_1(X0) | ! [X1] : (~v2_cat_1(X1) | ! [X2] : (~l1_cat_1(X2) | ~v2_cat_1(X2) | ! [X3] : (! [X4] : (r4_nattra_1(u1_cat_1(X0),u2_cat_1(X2),u1_cat_1(X0),u2_cat_1(X2),k6_isocat_1(X0,X1,X2,X3,X3,k7_nattra_1(X0,X1,X3),X4),k7_nattra_1(X0,X2,k2_isocat_1(X0,X1,X2,X3,X4))) | ~m2_cat_1(X4,X1,X2)) | ~m2_cat_1(X3,X0,X1))) | ~l1_cat_1(X1)) | ~v2_cat_1(X0)) [flattening 22625]
22918. ! [X1,X4,X0,X2,X5,X3] : ((r4_nattra_1(X4,X3,X1,X0,X5,X2) | ~r4_nattra_1(X4,X3,X1,X0,X2,X5)) | (~m1_relset_1(X2,X4,X3) | ~v1_funct_1(X5) | v1_xboole_0(X3) | v1_xboole_0(X0) | ~m1_relset_1(X5,X1,X0) | ~v1_funct_2(X2,X4,X3) | v1_xboole_0(X1) | ~v1_funct_1(X2) | v1_xboole_0(X4) | ~v1_funct_2(X5,X1,X0))) [ennf transformation 12844]
22919. ! [X2,X0,X3,X1,X4,X5] : (v1_xboole_0(X0) | v1_xboole_0(X3) | ~m1_relset_1(X5,X1,X0) | ~r4_nattra_1(X4,X3,X1,X0,X2,X5) | ~v1_funct_2(X2,X4,X3) | v1_xboole_0(X4) | ~v1_funct_1(X5) | ~m1_relset_1(X2,X4,X3) | ~v1_funct_2(X5,X1,X0) | r4_nattra_1(X4,X3,X1,X0,X5,X2) | ~v1_funct_1(X2) | v1_xboole_0(X1)) [flattening 22918]
23131. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (k11_isocat_2(X0,X1,X2,X3) = k2_isocat_1(X0,k11_cat_2(X1,X2),X1,X3,k8_isocat_2(X1,X2)) | ~m2_cat_1(X3,X0,k11_cat_2(X1,X2))) | (~v2_cat_1(X2) | ~l1_cat_1(X2))) | (~v2_cat_1(X1) | ~l1_cat_1(X1))) | (~v2_cat_1(X0) | ~l1_cat_1(X0))) [ennf transformation 11791]
23132. ! [X0] : (~l1_cat_1(X0) | ~v2_cat_1(X0) | ! [X1] : (! [X2] : (~l1_cat_1(X2) | ~v2_cat_1(X2) | ! [X3] : (k11_isocat_2(X0,X1,X2,X3) = k2_isocat_1(X0,k11_cat_2(X1,X2),X1,X3,k8_isocat_2(X1,X2)) | ~m2_cat_1(X3,X0,k11_cat_2(X1,X2)))) | ~l1_cat_1(X1) | ~v2_cat_1(X1))) [flattening 23131]
23213. ! [X0,X1,X2,X3] : (m2_cat_1(k12_isocat_2(X0,X2,X3,X1),X0,X3) | (~v2_cat_1(X2) | ~v2_cat_1(X3) | ~m2_cat_1(X1,X0,k11_cat_2(X2,X3)) | ~v2_cat_1(X0) | ~l1_cat_1(X0) | ~l1_cat_1(X3) | ~l1_cat_1(X2))) [ennf transformation 12827]
23214. ! [X2,X0,X3,X1] : (~m2_cat_1(X1,X0,k11_cat_2(X2,X3)) | ~v2_cat_1(X3) | ~v2_cat_1(X2) | ~l1_cat_1(X3) | ~v2_cat_1(X0) | ~l1_cat_1(X0) | m2_cat_1(k12_isocat_2(X0,X2,X3,X1),X0,X3) | ~l1_cat_1(X2)) [flattening 23213]
23455. ! [X1,X6,X3,X0,X5,X2,X4] : (m2_nattra_1(k6_isocat_1(X6,X0,X4,X2,X1,X3,X5),X6,X4,k2_isocat_1(X6,X0,X4,X2,X5),k2_isocat_1(X6,X0,X4,X1,X5)) | (~l1_cat_1(X6) | ~m2_nattra_1(X3,X6,X0,X2,X1) | ~v2_cat_1(X4) | ~v2_cat_1(X6) | ~l1_cat_1(X0) | ~m2_cat_1(X5,X0,X4) | ~v2_cat_1(X0) | ~l1_cat_1(X4) | ~m2_cat_1(X2,X6,X0) | ~m2_cat_1(X1,X6,X0))) [ennf transformation 13347]
23456. ! [X3,X6,X1,X0,X5,X2,X4] : (~v2_cat_1(X0) | ~l1_cat_1(X6) | ~l1_cat_1(X4) | ~v2_cat_1(X6) | ~m2_cat_1(X5,X0,X4) | ~v2_cat_1(X4) | ~m2_nattra_1(X3,X6,X0,X2,X1) | m2_nattra_1(k6_isocat_1(X6,X0,X4,X2,X1,X3,X5),X6,X4,k2_isocat_1(X6,X0,X4,X2,X5),k2_isocat_1(X6,X0,X4,X1,X5)) | ~l1_cat_1(X0) | ~m2_cat_1(X1,X6,X0) | ~m2_cat_1(X2,X6,X0)) [flattening 23455]
23749. ! [X1,X2,X0] : (m2_nattra_1(k7_nattra_1(X0,X2,X1),X0,X2,X1,X1) | (~m2_cat_1(X1,X0,X2) | ~l1_cat_1(X2) | ~l1_cat_1(X0) | ~v2_cat_1(X0) | ~v2_cat_1(X2))) [ennf transformation 14421]
23750. ! [X1,X0,X2] : (~l1_cat_1(X2) | ~v2_cat_1(X0) | ~m2_cat_1(X1,X0,X2) | ~v2_cat_1(X2) | m2_nattra_1(k7_nattra_1(X0,X2,X1),X0,X2,X1,X1) | ~l1_cat_1(X0)) [flattening 23749]
24631. ! [X0] : (~v1_xboole_0(u1_cat_1(X0)) | ~l1_cat_1(X0)) [ennf transformation 6631]
26379. ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((~r4_nattra_1(u1_cat_1(X0),u2_cat_1(X2),u1_cat_1(X0),u2_cat_1(X2),k7_nattra_1(X0,X2,k12_isocat_2(X0,X1,X2,X3)),k14_isocat_2(X0,X1,X2,X3,X3,k7_nattra_1(X0,k11_cat_2(X1,X2),X3))) | ~r4_nattra_1(u1_cat_1(X0),u2_cat_1(X1),u1_cat_1(X0),u2_cat_1(X1),k7_nattra_1(X0,X1,k11_isocat_2(X0,X1,X2,X3)),k13_isocat_2(X0,X1,X2,X3,X3,k7_nattra_1(X0,k11_cat_2(X1,X2),X3)))) & m2_cat_1(X3,X0,k11_cat_2(X1,X2))) & (l1_cat_1(X2) & v2_cat_1(X2))) & (l1_cat_1(X1) & v2_cat_1(X1))) & (l1_cat_1(X0) & v2_cat_1(X0))) [ennf transformation 11800]
26380. ? [X0] : (v2_cat_1(X0) & ? [X1] : (? [X2] : (v2_cat_1(X2) & ? [X3] : ((~r4_nattra_1(u1_cat_1(X0),u2_cat_1(X2),u1_cat_1(X0),u2_cat_1(X2),k7_nattra_1(X0,X2,k12_isocat_2(X0,X1,X2,X3)),k14_isocat_2(X0,X1,X2,X3,X3,k7_nattra_1(X0,k11_cat_2(X1,X2),X3))) | ~r4_nattra_1(u1_cat_1(X0),u2_cat_1(X1),u1_cat_1(X0),u2_cat_1(X1),k7_nattra_1(X0,X1,k11_isocat_2(X0,X1,X2,X3)),k13_isocat_2(X0,X1,X2,X3,X3,k7_nattra_1(X0,k11_cat_2(X1,X2),X3)))) & m2_cat_1(X3,X0,k11_cat_2(X1,X2))) & l1_cat_1(X2)) & l1_cat_1(X1) & v2_cat_1(X1)) & l1_cat_1(X0)) [flattening 26379]
27281. ! [X0] : (! [X1] : (! [X2] : (! [X3] : (k12_isocat_2(X0,X1,X2,X3) = k2_isocat_1(X0,k11_cat_2(X1,X2),X2,X3,k9_isocat_2(X1,X2)) | ~m2_cat_1(X3,X0,k11_cat_2(X1,X2))) | (~v2_cat_1(X2) | ~l1_cat_1(X2))) | (~v2_cat_1(X1) | ~l1_cat_1(X1))) | (~v2_cat_1(X0) | ~l1_cat_1(X0))) [ennf transformation 11792]
27282. ! [X0] : (~l1_cat_1(X0) | ~v2_cat_1(X0) | ! [X1] : (! [X2] : (! [X3] : (k12_isocat_2(X0,X1,X2,X3) = k2_isocat_1(X0,k11_cat_2(X1,X2),X2,X3,k9_isocat_2(X1,X2)) | ~m2_cat_1(X3,X0,k11_cat_2(X1,X2))) | ~l1_cat_1(X2) | ~v2_cat_1(X2)) | ~l1_cat_1(X1) | ~v2_cat_1(X1))) [flattening 27281]
28588. ! [X0] : (~l1_cat_1(X0) | ~v1_xboole_0(u2_cat_1(X0))) [ennf transformation 6632]
28806. ! [X1,X3,X0,X2] : (m2_cat_1(k11_isocat_2(X3,X1,X0,X2),X3,X1) | (~l1_cat_1(X1) | ~v2_cat_1(X0) | ~l1_cat_1(X3) | ~m2_cat_1(X2,X3,k11_cat_2(X1,X0)) | ~l1_cat_1(X0) | ~v2_cat_1(X1) | ~v2_cat_1(X3))) [ennf transformation 14159]
28807. ! [X3,X0,X1,X2] : (m2_cat_1(k11_isocat_2(X3,X1,X0,X2),X3,X1) | ~l1_cat_1(X0) | ~v2_cat_1(X0) | ~m2_cat_1(X2,X3,k11_cat_2(X1,X0)) | ~l1_cat_1(X3) | ~v2_cat_1(X1) | ~v2_cat_1(X3) | ~l1_cat_1(X1)) [flattening 28806]
31641. ! [X0,X1] : (~l1_cat_1(X0) | ~l1_cat_1(X1) | ~v2_cat_1(X0) | m2_cat_1(k17_cat_2(X0,X1),k11_cat_2(X0,X1),X1) | ~v2_cat_1(X1)) [rectify 19660]
31651. ! [X0,X1,X2,X3] : (m2_cat_1(k11_isocat_2(X0,X2,X1,X3),X0,X2) | ~l1_cat_1(X1) | ~v2_cat_1(X1) | ~m2_cat_1(X3,X0,k11_cat_2(X2,X1)) | ~l1_cat_1(X0) | ~v2_cat_1(X2) | ~v2_cat_1(X0) | ~l1_cat_1(X2)) [rectify 28807]
31697. ? [X0] : (v2_cat_1(X0) & ? [X1] : (? [X2] : (v2_cat_1(X2) & ? [X3] : ((~r4_nattra_1(u1_cat_1(X0),u2_cat_1(X2),u1_cat_1(X0),u2_cat_1(X2),k7_nattra_1(X0,X2,k12_isocat_2(X0,X1,X2,X3)),k14_isocat_2(X0,X1,X2,X3,X3,k7_nattra_1(X0,k11_cat_2(X1,X2),X3))) | ~r4_nattra_1(u1_cat_1(X0),u2_cat_1(X1),u1_cat_1(X0),u2_cat_1(X1),k7_nattra_1(X0,X1,k11_isocat_2(X0,X1,X2,X3)),k13_isocat_2(X0,X1,X2,X3,X3,k7_nattra_1(X0,k11_cat_2(X1,X2),X3)))) & m2_cat_1(X3,X0,k11_cat_2(X1,X2))) & l1_cat_1(X2)) & l1_cat_1(X1) & v2_cat_1(X1)) & l1_cat_1(X0)) => (v2_cat_1(sK817) & ? [X1] : (? [X2] : (v2_cat_1(X2) & ? [X3] : ((~r4_nattra_1(u1_cat_1(sK817),u2_cat_1(X2),u1_cat_1(sK817),u2_cat_1(X2),k7_nattra_1(sK817,X2,k12_isocat_2(sK817,X1,X2,X3)),k14_isocat_2(sK817,X1,X2,X3,X3,k7_nattra_1(sK817,k11_cat_2(X1,X2),X3))) | ~r4_nattra_1(u1_cat_1(sK817),u2_cat_1(X1),u1_cat_1(sK817),u2_cat_1(X1),k7_nattra_1(sK817,X1,k11_isocat_2(sK817,X1,X2,X3)),k13_isocat_2(sK817,X1,X2,X3,X3,k7_nattra_1(sK817,k11_cat_2(X1,X2),X3)))) & m2_cat_1(X3,sK817,k11_cat_2(X1,X2))) & l1_cat_1(X2)) & l1_cat_1(X1) & v2_cat_1(X1)) & l1_cat_1(sK817)) [choice axiom]
31698. ? [X1] : (? [X2] : (v2_cat_1(X2) & ? [X3] : ((~r4_nattra_1(u1_cat_1(sK817),u2_cat_1(X2),u1_cat_1(sK817),u2_cat_1(X2),k7_nattra_1(sK817,X2,k12_isocat_2(sK817,X1,X2,X3)),k14_isocat_2(sK817,X1,X2,X3,X3,k7_nattra_1(sK817,k11_cat_2(X1,X2),X3))) | ~r4_nattra_1(u1_cat_1(sK817),u2_cat_1(X1),u1_cat_1(sK817),u2_cat_1(X1),k7_nattra_1(sK817,X1,k11_isocat_2(sK817,X1,X2,X3)),k13_isocat_2(sK817,X1,X2,X3,X3,k7_nattra_1(sK817,k11_cat_2(X1,X2),X3)))) & m2_cat_1(X3,sK817,k11_cat_2(X1,X2))) & l1_cat_1(X2)) & l1_cat_1(X1) & v2_cat_1(X1)) => (? [X2] : (v2_cat_1(X2) & ? [X3] : ((~r4_nattra_1(u1_cat_1(sK817),u2_cat_1(X2),u1_cat_1(sK817),u2_cat_1(X2),k7_nattra_1(sK817,X2,k12_isocat_2(sK817,sK818,X2,X3)),k14_isocat_2(sK817,sK818,X2,X3,X3,k7_nattra_1(sK817,k11_cat_2(sK818,X2),X3))) | ~r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k7_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,X2,X3)),k13_isocat_2(sK817,sK818,X2,X3,X3,k7_nattra_1(sK817,k11_cat_2(sK818,X2),X3)))) & m2_cat_1(X3,sK817,k11_cat_2(sK818,X2))) & l1_cat_1(X2)) & l1_cat_1(sK818) & v2_cat_1(sK818)) [choice axiom]
31699. ? [X2] : (v2_cat_1(X2) & ? [X3] : ((~r4_nattra_1(u1_cat_1(sK817),u2_cat_1(X2),u1_cat_1(sK817),u2_cat_1(X2),k7_nattra_1(sK817,X2,k12_isocat_2(sK817,sK818,X2,X3)),k14_isocat_2(sK817,sK818,X2,X3,X3,k7_nattra_1(sK817,k11_cat_2(sK818,X2),X3))) | ~r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k7_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,X2,X3)),k13_isocat_2(sK817,sK818,X2,X3,X3,k7_nattra_1(sK817,k11_cat_2(sK818,X2),X3)))) & m2_cat_1(X3,sK817,k11_cat_2(sK818,X2))) & l1_cat_1(X2)) => (v2_cat_1(sK819) & ? [X3] : ((~r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK819),u1_cat_1(sK817),u2_cat_1(sK819),k7_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,X3)),k14_isocat_2(sK817,sK818,sK819,X3,X3,k7_nattra_1(sK817,k11_cat_2(sK818,sK819),X3))) | ~r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k7_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,X3)),k13_isocat_2(sK817,sK818,sK819,X3,X3,k7_nattra_1(sK817,k11_cat_2(sK818,sK819),X3)))) & m2_cat_1(X3,sK817,k11_cat_2(sK818,sK819))) & l1_cat_1(sK819)) [choice axiom]
31700. ? [X3] : ((~r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK819),u1_cat_1(sK817),u2_cat_1(sK819),k7_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,X3)),k14_isocat_2(sK817,sK818,sK819,X3,X3,k7_nattra_1(sK817,k11_cat_2(sK818,sK819),X3))) | ~r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k7_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,X3)),k13_isocat_2(sK817,sK818,sK819,X3,X3,k7_nattra_1(sK817,k11_cat_2(sK818,sK819),X3)))) & m2_cat_1(X3,sK817,k11_cat_2(sK818,sK819))) => ((~r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK819),u1_cat_1(sK817),u2_cat_1(sK819),k7_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k7_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | ~r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k7_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k7_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)))) & m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819))) [choice axiom]
31701. v2_cat_1(sK817) & ((v2_cat_1(sK819) & ((~r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK819),u1_cat_1(sK817),u2_cat_1(sK819),k7_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k7_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | ~r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k7_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k7_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)))) & m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819))) & l1_cat_1(sK819)) & l1_cat_1(sK818) & v2_cat_1(sK818)) & l1_cat_1(sK817) [skolemisation 26380,31700,31699,31698,31697]
34423. ! [X0,X1,X2,X3,X4,X5,X6] : (~v2_cat_1(X3) | ~l1_cat_1(X1) | ~l1_cat_1(X6) | ~v2_cat_1(X1) | ~m2_cat_1(X4,X3,X6) | ~v2_cat_1(X6) | ~m2_nattra_1(X0,X1,X3,X5,X2) | m2_nattra_1(k6_isocat_1(X1,X3,X6,X5,X2,X0,X4),X1,X6,k2_isocat_1(X1,X3,X6,X5,X4),k2_isocat_1(X1,X3,X6,X2,X4)) | ~l1_cat_1(X3) | ~m2_cat_1(X2,X1,X3) | ~m2_cat_1(X5,X1,X3)) [rectify 23456]
35511. ! [X0,X1,X2,X3] : (~m2_cat_1(X1,X2,X0) | ! [X4] : (~m1_nattra_1(X4,X2,X0,X1,X3) | (m2_relset_1(X4,u1_cat_1(X2),u2_cat_1(X0)) & v1_funct_1(X4) & v1_funct_2(X4,u1_cat_1(X2),u2_cat_1(X0)))) | ~l1_cat_1(X0) | ~l1_cat_1(X2) | ~v2_cat_1(X0) | ~m2_cat_1(X3,X2,X0) | ~v2_cat_1(X2)) [rectify 16934]
36222. ! [X0,X1,X2] : (~l1_cat_1(X2) | ~v2_cat_1(X1) | ~m2_cat_1(X0,X1,X2) | ~v2_cat_1(X2) | m2_nattra_1(k7_nattra_1(X1,X2,X0),X1,X2,X0,X0) | ~l1_cat_1(X1)) [rectify 23750]
36730. ! [X0,X1,X2,X3] : (~m2_cat_1(X3,X1,k11_cat_2(X0,X2)) | ~v2_cat_1(X2) | ~v2_cat_1(X0) | ~l1_cat_1(X2) | ~v2_cat_1(X1) | ~l1_cat_1(X1) | m2_cat_1(k12_isocat_2(X1,X0,X2,X3),X1,X2) | ~l1_cat_1(X0)) [rectify 23214]
37796. ! [X0,X1,X2,X3,X4,X5] : (~l1_cat_1(X0) | m2_nattra_1(k14_isocat_2(X2,X3,X0,X1,X4,X5),X2,X0,k12_isocat_2(X2,X3,X0,X1),k12_isocat_2(X2,X3,X0,X4)) | ~l1_cat_1(X3) | ~m2_cat_1(X1,X2,k11_cat_2(X3,X0)) | ~m2_cat_1(X4,X2,k11_cat_2(X3,X0)) | ~l1_cat_1(X2) | ~v2_cat_1(X2) | ~v2_cat_1(X0) | ~v2_cat_1(X3) | ~m2_nattra_1(X5,X2,k11_cat_2(X3,X0),X1,X4)) [rectify 20645]
38746. ! [X0,X1,X2,X3,X4,X5] : (v1_xboole_0(X1) | v1_xboole_0(X2) | ~m1_relset_1(X5,X3,X1) | ~r4_nattra_1(X4,X2,X3,X1,X0,X5) | ~v1_funct_2(X0,X4,X2) | v1_xboole_0(X4) | ~v1_funct_1(X5) | ~m1_relset_1(X0,X4,X2) | ~v1_funct_2(X5,X3,X1) | r4_nattra_1(X4,X2,X3,X1,X5,X0) | ~v1_funct_1(X0) | v1_xboole_0(X3)) [rectify 22919]
40573. ! [X0,X1,X2,X3] : (~l1_cat_1(X0) | ~m2_cat_1(X3,X0,X2) | ~v2_cat_1(X2) | ~m2_cat_1(X1,X0,X2) | ~v2_cat_1(X0) | ~l1_cat_1(X2) | ! [X4] : (m1_nattra_1(X4,X0,X2,X1,X3) | ~m2_nattra_1(X4,X0,X2,X1,X3))) [rectify 15397]
40629. ! [X0,X1] : (~l1_cat_1(X1) | ~v2_cat_1(X1) | ~l1_cat_1(X0) | (v2_cat_1(k11_cat_2(X1,X0)) & l1_cat_1(k11_cat_2(X1,X0))) | ~v2_cat_1(X0)) [rectify 16790]
41666. ! [X0,X1,X2,X3,X4,X5] : (~m2_nattra_1(X5,X0,k11_cat_2(X1,X2),X4,X3) | ~m2_cat_1(X4,X0,k11_cat_2(X1,X2)) | ~v2_cat_1(X2) | m2_nattra_1(k13_isocat_2(X0,X1,X2,X4,X3,X5),X0,X1,k11_isocat_2(X0,X1,X2,X4),k11_isocat_2(X0,X1,X2,X3)) | ~l1_cat_1(X1) | ~m2_cat_1(X3,X0,k11_cat_2(X1,X2)) | ~v2_cat_1(X1) | ~l1_cat_1(X0) | ~l1_cat_1(X2) | ~v2_cat_1(X0)) [rectify 22295]
42264. ! [X1,X2,X0] : ((m1_relset_1(X0,X2,X1) | ~m2_relset_1(X0,X2,X1)) & (m2_relset_1(X0,X2,X1) | ~m1_relset_1(X0,X2,X1))) [nnf transformation 12800]
42265. ! [X0,X1,X2] : ((m1_relset_1(X2,X1,X0) | ~m2_relset_1(X2,X1,X0)) & (m2_relset_1(X2,X1,X0) | ~m1_relset_1(X2,X1,X0))) [rectify 42264]
43842. ~m2_cat_1(X1,X0,X2) | ~l1_cat_1(X0) | ~v2_cat_1(X0) | ~l1_cat_1(X2) | k7_nattra_1(X0,X2,X1) = k4_nattra_1(X0,X2,X1) | ~v2_cat_1(X2) [cnf transformation 22275]
44441. m2_cat_1(k17_cat_2(X0,X1),k11_cat_2(X0,X1),X1) | ~l1_cat_1(X0) | ~v2_cat_1(X0) | ~l1_cat_1(X1) | ~v2_cat_1(X1) [cnf transformation 31641]
44450. m2_cat_1(k11_isocat_2(X0,X2,X1,X3),X0,X2) | ~v2_cat_1(X0) | ~v2_cat_1(X1) | ~v2_cat_1(X2) | ~m2_cat_1(X3,X0,k11_cat_2(X2,X1)) | ~l1_cat_1(X0) | ~l1_cat_1(X1) | ~l1_cat_1(X2) [cnf transformation 31651]
44562. l1_cat_1(sK817) [cnf transformation 31701]
44563. v2_cat_1(sK818) [cnf transformation 31701]
44564. l1_cat_1(sK818) [cnf transformation 31701]
44565. l1_cat_1(sK819) [cnf transformation 31701]
44566. m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) [cnf transformation 31701]
44567. ~r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK819),u1_cat_1(sK817),u2_cat_1(sK819),k7_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k7_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | ~r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k7_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k7_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) [cnf transformation 31701]
44568. v2_cat_1(sK819) [cnf transformation 31701]
44569. v2_cat_1(sK817) [cnf transformation 31701]
46878. r4_nattra_1(u1_cat_1(X0),u2_cat_1(X2),u1_cat_1(X0),u2_cat_1(X2),k6_isocat_1(X0,X1,X2,X3,X3,k7_nattra_1(X0,X1,X3),X4),k7_nattra_1(X0,X2,k2_isocat_1(X0,X1,X2,X3,X4))) | ~l1_cat_1(X0) | ~m2_cat_1(X4,X1,X2) | ~m2_cat_1(X3,X0,X1) | ~l1_cat_1(X2) | ~v2_cat_1(X0) | ~v2_cat_1(X1) | ~v2_cat_1(X2) | ~l1_cat_1(X1) [cnf transformation 22626]
48032. ~v2_cat_1(X1) | ~l1_cat_1(X1) | k17_cat_2(X1,X0) = k9_isocat_2(X1,X0) | ~l1_cat_1(X0) | ~v2_cat_1(X0) [cnf transformation 15341]
48525. ~v1_xboole_0(u2_cat_1(X0)) | ~l1_cat_1(X0) [cnf transformation 28588]
49971. m2_nattra_1(k6_isocat_1(X1,X3,X6,X5,X2,X0,X4),X1,X6,k2_isocat_1(X1,X3,X6,X5,X4),k2_isocat_1(X1,X3,X6,X2,X4)) | ~v2_cat_1(X1) | ~l1_cat_1(X3) | ~m2_cat_1(X2,X1,X3) | ~l1_cat_1(X1) | ~l1_cat_1(X6) | ~m2_cat_1(X5,X1,X3) | ~m2_nattra_1(X0,X1,X3,X5,X2) | ~v2_cat_1(X6) | ~v2_cat_1(X3) | ~m2_cat_1(X4,X3,X6) [cnf transformation 34423]
51988. ~v2_cat_1(X1) | ~v2_cat_1(X0) | k8_isocat_2(X1,X0) = k16_cat_2(X1,X0) | ~l1_cat_1(X0) | ~l1_cat_1(X1) [cnf transformation 21276]
52011. ~m1_nattra_1(X4,X2,X0,X1,X3) | ~l1_cat_1(X2) | ~v2_cat_1(X2) | ~m2_cat_1(X3,X2,X0) | ~l1_cat_1(X0) | ~m2_cat_1(X1,X2,X0) | v1_funct_2(X4,u1_cat_1(X2),u2_cat_1(X0)) | ~v2_cat_1(X0) [cnf transformation 35511]
52012. ~m1_nattra_1(X4,X2,X0,X1,X3) | ~v2_cat_1(X0) | v1_funct_1(X4) | ~m2_cat_1(X1,X2,X0) | ~l1_cat_1(X0) | ~m2_cat_1(X3,X2,X0) | ~l1_cat_1(X2) | ~v2_cat_1(X2) [cnf transformation 35511]
52013. ~m1_nattra_1(X4,X2,X0,X1,X3) | ~m2_cat_1(X3,X2,X0) | ~l1_cat_1(X2) | m2_relset_1(X4,u1_cat_1(X2),u2_cat_1(X0)) | ~l1_cat_1(X0) | ~v2_cat_1(X2) | ~v2_cat_1(X0) | ~m2_cat_1(X1,X2,X0) [cnf transformation 35511]
53064. m2_cat_1(k8_isocat_2(X0,X1),k11_cat_2(X0,X1),X0) | ~l1_cat_1(X1) | ~l1_cat_1(X0) | ~v2_cat_1(X1) | ~v2_cat_1(X0) [cnf transformation 14877]
53438. m2_nattra_1(k7_nattra_1(X1,X2,X0),X1,X2,X0,X0) | ~m2_cat_1(X0,X1,X2) | ~l1_cat_1(X1) | ~l1_cat_1(X2) | ~v2_cat_1(X1) | ~v2_cat_1(X2) [cnf transformation 36222]
54428. m2_cat_1(k12_isocat_2(X1,X0,X2,X3),X1,X2) | ~l1_cat_1(X0) | ~l1_cat_1(X2) | ~l1_cat_1(X1) | ~m2_cat_1(X3,X1,k11_cat_2(X0,X2)) | ~v2_cat_1(X1) | ~v2_cat_1(X2) | ~v2_cat_1(X0) [cnf transformation 36730]
55258. ~m2_nattra_1(X5,X0,k11_cat_2(X1,X2),X3,X4) | ~v2_cat_1(X2) | ~l1_cat_1(X1) | ~l1_cat_1(X0) | ~m2_cat_1(X3,X0,k11_cat_2(X1,X2)) | ~m2_cat_1(X4,X0,k11_cat_2(X1,X2)) | ~v2_cat_1(X1) | k14_isocat_2(X0,X1,X2,X3,X4,X5) = k6_isocat_1(X0,k11_cat_2(X1,X2),X2,X3,X4,X5,k9_isocat_2(X1,X2)) | ~l1_cat_1(X2) | ~v2_cat_1(X0) [cnf transformation 17430]
56600. m2_nattra_1(k14_isocat_2(X2,X3,X0,X1,X4,X5),X2,X0,k12_isocat_2(X2,X3,X0,X1),k12_isocat_2(X2,X3,X0,X4)) | ~m2_nattra_1(X5,X2,k11_cat_2(X3,X0),X1,X4) | ~v2_cat_1(X2) | ~m2_cat_1(X4,X2,k11_cat_2(X3,X0)) | ~m2_cat_1(X1,X2,k11_cat_2(X3,X0)) | ~v2_cat_1(X3) | ~l1_cat_1(X2) | ~l1_cat_1(X3) | ~l1_cat_1(X0) | ~v2_cat_1(X0) [cnf transformation 37796]
58555. ~r4_nattra_1(X4,X2,X3,X1,X0,X5) | v1_xboole_0(X1) | ~m1_relset_1(X0,X4,X2) | ~m1_relset_1(X5,X3,X1) | ~v1_funct_1(X0) | r4_nattra_1(X4,X2,X3,X1,X5,X0) | ~v1_funct_2(X0,X4,X2) | v1_xboole_0(X4) | v1_xboole_0(X2) | ~v1_funct_1(X5) | v1_xboole_0(X3) | ~v1_funct_2(X5,X3,X1) [cnf transformation 38746]
59071. ~m2_cat_1(X3,X0,k11_cat_2(X1,X2)) | ~l1_cat_1(X1) | ~l1_cat_1(X2) | k11_isocat_2(X0,X1,X2,X3) = k2_isocat_1(X0,k11_cat_2(X1,X2),X1,X3,k8_isocat_2(X1,X2)) | ~v2_cat_1(X0) | ~l1_cat_1(X0) | ~v2_cat_1(X1) | ~v2_cat_1(X2) [cnf transformation 23132]
59978. ~m2_nattra_1(X5,X0,k11_cat_2(X1,X2),X3,X4) | ~l1_cat_1(X2) | ~l1_cat_1(X0) | ~v2_cat_1(X1) | ~m2_cat_1(X4,X0,k11_cat_2(X1,X2)) | k13_isocat_2(X0,X1,X2,X3,X4,X5) = k6_isocat_1(X0,k11_cat_2(X1,X2),X1,X3,X4,X5,k8_isocat_2(X1,X2)) | ~v2_cat_1(X0) | ~v2_cat_1(X2) | ~m2_cat_1(X3,X0,k11_cat_2(X1,X2)) | ~l1_cat_1(X1) [cnf transformation 17306]
62059. ~m2_nattra_1(X4,X0,X2,X1,X3) | m1_nattra_1(X4,X0,X2,X1,X3) | ~v2_cat_1(X0) | ~m2_cat_1(X3,X0,X2) | ~l1_cat_1(X0) | ~m2_cat_1(X1,X0,X2) | ~v2_cat_1(X2) | ~l1_cat_1(X2) [cnf transformation 40573]
62120. ~m2_cat_1(X3,X0,k11_cat_2(X1,X2)) | ~l1_cat_1(X0) | ~v2_cat_1(X1) | ~v2_cat_1(X2) | ~v2_cat_1(X0) | ~l1_cat_1(X2) | ~l1_cat_1(X1) | k12_isocat_2(X0,X1,X2,X3) = k2_isocat_1(X0,k11_cat_2(X1,X2),X2,X3,k9_isocat_2(X1,X2)) [cnf transformation 27282]
62182. l1_cat_1(k11_cat_2(X1,X0)) | ~v2_cat_1(X1) | ~l1_cat_1(X1) | ~l1_cat_1(X0) | ~v2_cat_1(X0) [cnf transformation 40629]
62183. v2_cat_1(k11_cat_2(X1,X0)) | ~l1_cat_1(X0) | ~v2_cat_1(X1) | ~l1_cat_1(X1) | ~v2_cat_1(X0) [cnf transformation 40629]
63394. ~v1_xboole_0(u1_cat_1(X0)) | ~l1_cat_1(X0) [cnf transformation 24631]
64143. m2_nattra_1(k13_isocat_2(X0,X1,X2,X4,X3,X5),X0,X1,k11_isocat_2(X0,X1,X2,X4),k11_isocat_2(X0,X1,X2,X3)) | ~m2_nattra_1(X5,X0,k11_cat_2(X1,X2),X4,X3) | ~m2_cat_1(X3,X0,k11_cat_2(X1,X2)) | ~l1_cat_1(X0) | ~l1_cat_1(X1) | ~v2_cat_1(X0) | ~m2_cat_1(X4,X0,k11_cat_2(X1,X2)) | ~l1_cat_1(X2) | ~v2_cat_1(X1) | ~v2_cat_1(X2) [cnf transformation 41666]
65267. ~m2_relset_1(X2,X1,X0) | m1_relset_1(X2,X1,X0) [cnf transformation 42265]
78352. 369 <=> r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK819),u1_cat_1(sK817),u2_cat_1(sK819),k7_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k7_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) [avatar definition]
78354. ~r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK819),u1_cat_1(sK817),u2_cat_1(sK819),k7_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k7_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) <- (~369) [avatar component clause 78352]
78356. 370 <=> r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k7_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k7_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) [avatar definition]
78358. ~r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k7_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k7_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) <- (~370) [avatar component clause 78356]
78359. ~369 | ~370 [avatar split clause 44567,78356,78352]
140443. 1515 <=> v2_cat_1(k11_cat_2(sK818,sK819)) [avatar definition]
140444. v2_cat_1(k11_cat_2(sK818,sK819)) <- (1515) [avatar component clause 140443]
140445. ~v2_cat_1(k11_cat_2(sK818,sK819)) <- (~1515) [avatar component clause 140443]
140451. 1517 <=> l1_cat_1(k11_cat_2(sK818,sK819)) [avatar definition]
140452. l1_cat_1(k11_cat_2(sK818,sK819)) <- (1517) [avatar component clause 140451]
140453. ~l1_cat_1(k11_cat_2(sK818,sK819)) <- (~1517) [avatar component clause 140451]
140472. ~m2_cat_1(X0,X1,X2) | ~l1_cat_1(X2) | ~v2_cat_1(X1) | ~l1_cat_1(X1) | ~l1_cat_1(X1) | ~v2_cat_1(X1) | ~v2_cat_1(X2) | ~m2_cat_1(X0,X1,X2) | ~l1_cat_1(X2) | m1_nattra_1(k7_nattra_1(X1,X2,X0),X1,X2,X0,X0) | ~v2_cat_1(X2) | ~m2_cat_1(X0,X1,X2) [resolution 53438,62059]
140473. m1_nattra_1(k7_nattra_1(X1,X2,X0),X1,X2,X0,X0) | ~v2_cat_1(X1) | ~l1_cat_1(X1) | ~v2_cat_1(X2) | ~l1_cat_1(X2) | ~m2_cat_1(X0,X1,X2) [duplicate literal removal 140472]
140474. ~l1_cat_1(sK819) | ~v2_cat_1(sK819) | ~v2_cat_1(sK818) | ~l1_cat_1(sK818) <- (~1515) [resolution 62183,140445]
140475. ~l1_cat_1(sK818) | ~v2_cat_1(sK818) | ~v2_cat_1(sK819) <- (~1515) [subsumption resolution 140474,44565]
140476. ~v2_cat_1(sK819) | ~v2_cat_1(sK818) <- (~1515) [subsumption resolution 140475,44564]
140477. ~v2_cat_1(sK818) <- (~1515) [subsumption resolution 140476,44568]
140478. $false <- (~1515) [subsumption resolution 140477,44563]
140479. 1515 [avatar contradiction clause 140478]
140488. ~l1_cat_1(sK819) | ~v2_cat_1(sK819) | ~l1_cat_1(sK818) | ~v2_cat_1(sK818) <- (~1517) [resolution 62182,140453]
140489. ~v2_cat_1(sK818) | ~l1_cat_1(sK818) | ~v2_cat_1(sK819) <- (~1517) [subsumption resolution 140488,44565]
140490. ~v2_cat_1(sK819) | ~l1_cat_1(sK818) <- (~1517) [subsumption resolution 140489,44563]
140491. ~l1_cat_1(sK818) <- (~1517) [subsumption resolution 140490,44568]
140492. $false <- (~1517) [subsumption resolution 140491,44564]
140493. 1517 [avatar contradiction clause 140492]
140496. ~l1_cat_1(X2) | ~v2_cat_1(X2) | ~m2_cat_1(X4,X1,k11_cat_2(X2,X3)) | ~l1_cat_1(X3) | ~v2_cat_1(X1) | ~v2_cat_1(X3) | ~l1_cat_1(X2) | ~l1_cat_1(X1) | ~m2_cat_1(X5,X1,k11_cat_2(X2,X3)) | ~v2_cat_1(X2) | ~m2_cat_1(k11_isocat_2(X1,X2,X3,X4),X1,X2) | m1_nattra_1(k13_isocat_2(X1,X2,X3,X4,X5,X0),X1,X2,k11_isocat_2(X1,X2,X3,X4),k11_isocat_2(X1,X2,X3,X5)) | ~v2_cat_1(X1) | ~m2_cat_1(k11_isocat_2(X1,X2,X3,X5),X1,X2) | ~m2_nattra_1(X0,X1,k11_cat_2(X2,X3),X4,X5) | ~l1_cat_1(X1) [resolution 64143,62059]
140497. ~v2_cat_1(X2) | ~m2_cat_1(X5,X1,k11_cat_2(X2,X3)) | ~l1_cat_1(X2) | ~m2_cat_1(k11_isocat_2(X1,X2,X3,X4),X1,X2) | ~m2_nattra_1(X0,X1,k11_cat_2(X2,X3),X4,X5) | ~l1_cat_1(X1) | ~l1_cat_1(X3) | ~v2_cat_1(X1) | ~m2_cat_1(k11_isocat_2(X1,X2,X3,X5),X1,X2) | ~v2_cat_1(X3) | m1_nattra_1(k13_isocat_2(X1,X2,X3,X4,X5,X0),X1,X2,k11_isocat_2(X1,X2,X3,X4),k11_isocat_2(X1,X2,X3,X5)) | ~m2_cat_1(X4,X1,k11_cat_2(X2,X3)) [duplicate literal removal 140496]
140498. ~m2_cat_1(X4,X1,k11_cat_2(X2,X3)) | ~l1_cat_1(X1) | ~v2_cat_1(X3) | ~v2_cat_1(X2) | ~m2_cat_1(X5,X1,k11_cat_2(X2,X3)) | ~l1_cat_1(X3) | ~m2_cat_1(k11_isocat_2(X1,X2,X3,X5),X1,X2) | ~m2_nattra_1(X0,X1,k11_cat_2(X2,X3),X4,X5) | m1_nattra_1(k13_isocat_2(X1,X2,X3,X4,X5,X0),X1,X2,k11_isocat_2(X1,X2,X3,X4),k11_isocat_2(X1,X2,X3,X5)) | ~l1_cat_1(X2) | ~v2_cat_1(X1) [subsumption resolution 140497,44450]
140499. m1_nattra_1(k13_isocat_2(X1,X2,X3,X4,X5,X0),X1,X2,k11_isocat_2(X1,X2,X3,X4),k11_isocat_2(X1,X2,X3,X5)) | ~v2_cat_1(X1) | ~v2_cat_1(X3) | ~l1_cat_1(X2) | ~m2_cat_1(X5,X1,k11_cat_2(X2,X3)) | ~m2_nattra_1(X0,X1,k11_cat_2(X2,X3),X4,X5) | ~v2_cat_1(X2) | ~l1_cat_1(X1) | ~l1_cat_1(X3) | ~m2_cat_1(X4,X1,k11_cat_2(X2,X3)) [subsumption resolution 140498,44450]
140500. ~m2_cat_1(k12_isocat_2(X1,X2,X3,X5),X1,X3) | ~l1_cat_1(X2) | ~v2_cat_1(X2) | ~m2_nattra_1(X0,X1,k11_cat_2(X2,X3),X4,X5) | ~l1_cat_1(X3) | ~l1_cat_1(X3) | m1_nattra_1(k14_isocat_2(X1,X2,X3,X4,X5,X0),X1,X3,k12_isocat_2(X1,X2,X3,X4),k12_isocat_2(X1,X2,X3,X5)) | ~m2_cat_1(X4,X1,k11_cat_2(X2,X3)) | ~v2_cat_1(X3) | ~m2_cat_1(k12_isocat_2(X1,X2,X3,X4),X1,X3) | ~l1_cat_1(X1) | ~v2_cat_1(X3) | ~l1_cat_1(X1) | ~v2_cat_1(X1) | ~m2_cat_1(X5,X1,k11_cat_2(X2,X3)) | ~v2_cat_1(X1) [resolution 56600,62059]
140501. ~m2_cat_1(k12_isocat_2(X1,X2,X3,X4),X1,X3) | ~m2_cat_1(k12_isocat_2(X1,X2,X3,X5),X1,X3) | m1_nattra_1(k14_isocat_2(X1,X2,X3,X4,X5,X0),X1,X3,k12_isocat_2(X1,X2,X3,X4),k12_isocat_2(X1,X2,X3,X5)) | ~m2_cat_1(X5,X1,k11_cat_2(X2,X3)) | ~l1_cat_1(X2) | ~m2_cat_1(X4,X1,k11_cat_2(X2,X3)) | ~l1_cat_1(X1) | ~v2_cat_1(X2) | ~v2_cat_1(X1) | ~v2_cat_1(X3) | ~l1_cat_1(X3) | ~m2_nattra_1(X0,X1,k11_cat_2(X2,X3),X4,X5) [duplicate literal removal 140500]
140502. ~l1_cat_1(X2) | ~v2_cat_1(X1) | ~v2_cat_1(X3) | ~m2_cat_1(X4,X1,k11_cat_2(X2,X3)) | ~m2_cat_1(X5,X1,k11_cat_2(X2,X3)) | ~v2_cat_1(X2) | ~m2_cat_1(k12_isocat_2(X1,X2,X3,X4),X1,X3) | ~m2_nattra_1(X0,X1,k11_cat_2(X2,X3),X4,X5) | ~l1_cat_1(X1) | ~l1_cat_1(X3) | m1_nattra_1(k14_isocat_2(X1,X2,X3,X4,X5,X0),X1,X3,k12_isocat_2(X1,X2,X3,X4),k12_isocat_2(X1,X2,X3,X5)) [subsumption resolution 140501,54428]
140503. m1_nattra_1(k14_isocat_2(X1,X2,X3,X4,X5,X0),X1,X3,k12_isocat_2(X1,X2,X3,X4),k12_isocat_2(X1,X2,X3,X5)) | ~l1_cat_1(X3) | ~v2_cat_1(X1) | ~v2_cat_1(X3) | ~m2_cat_1(X5,X1,k11_cat_2(X2,X3)) | ~l1_cat_1(X1) | ~m2_cat_1(X4,X1,k11_cat_2(X2,X3)) | ~l1_cat_1(X2) | ~m2_nattra_1(X0,X1,k11_cat_2(X2,X3),X4,X5) | ~v2_cat_1(X2) [subsumption resolution 140502,54428]
140514. ~v2_cat_1(X0) | ~v2_cat_1(X1) | ~m2_cat_1(X3,X0,k11_cat_2(X1,X2)) | ~l1_cat_1(X0) | ~l1_cat_1(X1) | ~l1_cat_1(X1) | ~v2_cat_1(X2) | ~v2_cat_1(X1) | ~l1_cat_1(X0) | ~v2_cat_1(X0) | k7_nattra_1(X0,X1,k11_isocat_2(X0,X1,X2,X3)) = k4_nattra_1(X0,X1,k11_isocat_2(X0,X1,X2,X3)) | ~l1_cat_1(X2) [resolution 43842,44450]
140515. ~l1_cat_1(X2) | ~v2_cat_1(X0) | ~l1_cat_1(X0) | ~l1_cat_1(X1) | ~v2_cat_1(X2) | ~v2_cat_1(X0) | ~v2_cat_1(X1) | k7_nattra_1(X0,X1,k12_isocat_2(X0,X2,X1,X3)) = k4_nattra_1(X0,X1,k12_isocat_2(X0,X2,X1,X3)) | ~m2_cat_1(X3,X0,k11_cat_2(X2,X1)) | ~l1_cat_1(X1) | ~v2_cat_1(X1) | ~l1_cat_1(X0) [resolution 43842,54428]
140516. ~l1_cat_1(k11_cat_2(sK818,sK819)) | ~l1_cat_1(sK817) | ~v2_cat_1(sK817) | k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820) = k7_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820) | ~v2_cat_1(k11_cat_2(sK818,sK819)) [resolution 43842,44566]
140518. ~m2_cat_1(X3,X0,k11_cat_2(X2,X1)) | k7_nattra_1(X0,X1,k12_isocat_2(X0,X2,X1,X3)) = k4_nattra_1(X0,X1,k12_isocat_2(X0,X2,X1,X3)) | ~v2_cat_1(X2) | ~l1_cat_1(X0) | ~v2_cat_1(X1) | ~l1_cat_1(X2) | ~v2_cat_1(X0) | ~l1_cat_1(X1) [duplicate literal removal 140515]
140520. ~m2_cat_1(X3,X0,k11_cat_2(X1,X2)) | ~l1_cat_1(X2) | ~v2_cat_1(X2) | ~v2_cat_1(X1) | ~v2_cat_1(X0) | k7_nattra_1(X0,X1,k11_isocat_2(X0,X1,X2,X3)) = k4_nattra_1(X0,X1,k11_isocat_2(X0,X1,X2,X3)) | ~l1_cat_1(X0) | ~l1_cat_1(X1) [duplicate literal removal 140514]
140521. ~v2_cat_1(k11_cat_2(sK818,sK819)) | ~v2_cat_1(sK817) | k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820) = k7_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820) | ~l1_cat_1(sK817) <- (1517) [subsumption resolution 140516,140452]
140524. ~l1_cat_1(sK817) | ~v2_cat_1(sK817) | k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820) = k7_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820) <- (1515, 1517) [subsumption resolution 140521,140444]
140527. k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820) = k7_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820) | ~v2_cat_1(sK817) <- (1515, 1517) [subsumption resolution 140524,44562]
140528. k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820) = k7_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820) <- (1515, 1517) [subsumption resolution 140527,44569]
140529. ~r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k7_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) <- (~370, 1515, 1517) [backward demodulation 78358,140528]
140530. ~l1_cat_1(k11_cat_2(sK818,sK819)) | m2_nattra_1(k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),sK817,k11_cat_2(sK818,sK819),sK820,sK820) | ~l1_cat_1(sK817) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~v2_cat_1(k11_cat_2(sK818,sK819)) | ~v2_cat_1(sK817) <- (1515, 1517) [superposition 53438,140528]
140531. ~l1_cat_1(sK817) | ~v2_cat_1(k11_cat_2(sK818,sK819)) | ~v2_cat_1(sK817) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | m2_nattra_1(k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),sK817,k11_cat_2(sK818,sK819),sK820,sK820) <- (1515, 1517) [subsumption resolution 140530,140452]
140532. ~v2_cat_1(sK817) | ~v2_cat_1(k11_cat_2(sK818,sK819)) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | m2_nattra_1(k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),sK817,k11_cat_2(sK818,sK819),sK820,sK820) <- (1515, 1517) [subsumption resolution 140531,44562]
140533. m2_nattra_1(k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),sK817,k11_cat_2(sK818,sK819),sK820,sK820) | ~v2_cat_1(k11_cat_2(sK818,sK819)) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) <- (1515, 1517) [subsumption resolution 140532,44569]
140534. m2_nattra_1(k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),sK817,k11_cat_2(sK818,sK819),sK820,sK820) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) <- (1515, 1517) [subsumption resolution 140533,140444]
140535. m2_nattra_1(k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),sK817,k11_cat_2(sK818,sK819),sK820,sK820) <- (1515, 1517) [subsumption resolution 140534,44566]
140542. k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)) = k7_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)) | ~l1_cat_1(sK818) | ~v2_cat_1(sK819) | ~v2_cat_1(sK818) | ~l1_cat_1(sK819) | ~v2_cat_1(sK817) | ~l1_cat_1(sK817) [resolution 140520,44566]
140547. k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)) = k7_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)) | ~v2_cat_1(sK819) | ~v2_cat_1(sK818) | ~l1_cat_1(sK817) | ~v2_cat_1(sK817) | ~l1_cat_1(sK819) [subsumption resolution 140542,44564]
140552. ~v2_cat_1(sK818) | k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)) = k7_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)) | ~l1_cat_1(sK819) | ~l1_cat_1(sK817) | ~v2_cat_1(sK817) [subsumption resolution 140547,44568]
140556. ~l1_cat_1(sK817) | ~v2_cat_1(sK817) | k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)) = k7_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)) | ~l1_cat_1(sK819) [subsumption resolution 140552,44563]
140558. k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)) = k7_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)) | ~v2_cat_1(sK817) | ~l1_cat_1(sK819) [subsumption resolution 140556,44562]
140559. ~l1_cat_1(sK819) | k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)) = k7_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)) [subsumption resolution 140558,44569]
140560. k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)) = k7_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)) [subsumption resolution 140559,44565]
140561. ~r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) <- (~370, 1515, 1517) [backward demodulation 140529,140560]
140572. 1521 <=> m2_cat_1(k11_isocat_2(sK817,sK818,sK819,sK820),sK817,sK818) [avatar definition]
140573. m2_cat_1(k11_isocat_2(sK817,sK818,sK819,sK820),sK817,sK818) <- (1521) [avatar component clause 140572]
140574. ~m2_cat_1(k11_isocat_2(sK817,sK818,sK819,sK820),sK817,sK818) <- (~1521) [avatar component clause 140572]
140580. ~l1_cat_1(sK819) | k7_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)) = k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)) | ~v2_cat_1(sK818) | ~l1_cat_1(sK818) | ~l1_cat_1(sK817) | ~v2_cat_1(sK817) | ~v2_cat_1(sK819) [resolution 140518,44566]
140584. k7_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)) = k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)) | ~v2_cat_1(sK819) | ~l1_cat_1(sK817) | ~v2_cat_1(sK817) | ~v2_cat_1(sK818) | ~l1_cat_1(sK818) [subsumption resolution 140580,44565]
140589. ~l1_cat_1(sK817) | ~v2_cat_1(sK817) | ~v2_cat_1(sK818) | k7_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)) = k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)) | ~l1_cat_1(sK818) [subsumption resolution 140584,44568]
140593. k7_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)) = k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)) | ~v2_cat_1(sK817) | ~v2_cat_1(sK818) | ~l1_cat_1(sK818) [subsumption resolution 140589,44562]
140594. k7_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)) = k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)) | ~v2_cat_1(sK818) | ~l1_cat_1(sK818) [subsumption resolution 140593,44569]
140595. k7_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)) = k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)) | ~l1_cat_1(sK818) [subsumption resolution 140594,44563]
140596. k7_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)) = k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)) [subsumption resolution 140595,44564]
140607. 1523 <=> m2_cat_1(k12_isocat_2(sK817,sK818,sK819,sK820),sK817,sK819) [avatar definition]
140608. m2_cat_1(k12_isocat_2(sK817,sK818,sK819,sK820),sK817,sK819) <- (1523) [avatar component clause 140607]
140609. ~m2_cat_1(k12_isocat_2(sK817,sK818,sK819,sK820),sK817,sK819) <- (~1523) [avatar component clause 140607]
140611. ~v2_cat_1(sK819) | ~l1_cat_1(sK819) | ~v2_cat_1(sK817) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~v2_cat_1(sK818) | ~l1_cat_1(sK818) | ~l1_cat_1(sK817) <- (~1521) [resolution 140574,44450]
140612. ~l1_cat_1(sK817) | ~l1_cat_1(sK818) | ~l1_cat_1(sK819) | ~v2_cat_1(sK817) | ~v2_cat_1(sK818) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) <- (~1521) [subsumption resolution 140611,44568]
140613. ~l1_cat_1(sK818) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~v2_cat_1(sK818) | ~v2_cat_1(sK817) | ~l1_cat_1(sK819) <- (~1521) [subsumption resolution 140612,44562]
140614. ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~v2_cat_1(sK817) | ~l1_cat_1(sK819) | ~v2_cat_1(sK818) <- (~1521) [subsumption resolution 140613,44564]
140615. ~v2_cat_1(sK817) | ~v2_cat_1(sK818) | ~l1_cat_1(sK819) <- (~1521) [subsumption resolution 140614,44566]
140616. ~l1_cat_1(sK819) | ~v2_cat_1(sK818) <- (~1521) [subsumption resolution 140615,44569]
140617. ~v2_cat_1(sK818) <- (~1521) [subsumption resolution 140616,44565]
140618. $false <- (~1521) [subsumption resolution 140617,44563]
140619. 1521 [avatar contradiction clause 140618]
140623. ~l1_cat_1(sK817) | ~v2_cat_1(sK819) | k6_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),k8_isocat_2(sK818,sK819)) = k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)) | ~l1_cat_1(sK819) | ~v2_cat_1(sK818) | ~v2_cat_1(sK817) | ~l1_cat_1(sK818) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) <- (1515, 1517) [resolution 59978,140535]
140626. ~v2_cat_1(sK819) | ~l1_cat_1(sK818) | ~v2_cat_1(sK818) | ~l1_cat_1(sK819) | ~v2_cat_1(sK817) | k6_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),k8_isocat_2(sK818,sK819)) = k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)) | ~l1_cat_1(sK817) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) <- (1515, 1517) [duplicate literal removal 140623]
140630. ~v2_cat_1(sK818) | k6_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),k8_isocat_2(sK818,sK819)) = k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)) | ~l1_cat_1(sK817) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~v2_cat_1(sK817) | ~l1_cat_1(sK819) | ~l1_cat_1(sK818) <- (1515, 1517) [subsumption resolution 140626,44568]
140634. k6_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),k8_isocat_2(sK818,sK819)) = k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~l1_cat_1(sK817) | ~v2_cat_1(sK817) | ~l1_cat_1(sK818) | ~l1_cat_1(sK819) <- (1515, 1517) [subsumption resolution 140630,44563]
140638. ~l1_cat_1(sK819) | ~l1_cat_1(sK817) | k6_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),k8_isocat_2(sK818,sK819)) = k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)) | ~l1_cat_1(sK818) | ~v2_cat_1(sK817) <- (1515, 1517) [subsumption resolution 140634,44566]
140641. k6_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),k8_isocat_2(sK818,sK819)) = k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)) | ~l1_cat_1(sK818) | ~l1_cat_1(sK817) | ~v2_cat_1(sK817) <- (1515, 1517) [subsumption resolution 140638,44565]
140642. k6_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),k8_isocat_2(sK818,sK819)) = k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)) | ~l1_cat_1(sK817) | ~v2_cat_1(sK817) <- (1515, 1517) [subsumption resolution 140641,44564]
140643. ~v2_cat_1(sK817) | k6_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),k8_isocat_2(sK818,sK819)) = k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)) <- (1515, 1517) [subsumption resolution 140642,44562]
140644. k6_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),k8_isocat_2(sK818,sK819)) = k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)) <- (1515, 1517) [subsumption resolution 140643,44569]
140666. ~l1_cat_1(sK817) | ~v2_cat_1(sK817) | ~l1_cat_1(sK819) | ~l1_cat_1(sK818) | ~v2_cat_1(sK818) | ~v2_cat_1(sK819) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) <- (~1523) [resolution 140609,54428]
140667. ~v2_cat_1(sK818) | ~l1_cat_1(sK819) | ~v2_cat_1(sK819) | ~v2_cat_1(sK817) | ~l1_cat_1(sK818) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) <- (~1523) [subsumption resolution 140666,44562]
140668. ~l1_cat_1(sK819) | ~v2_cat_1(sK819) | ~v2_cat_1(sK817) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~l1_cat_1(sK818) <- (~1523) [subsumption resolution 140667,44563]
140669. ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~l1_cat_1(sK818) | ~v2_cat_1(sK819) | ~v2_cat_1(sK817) <- (~1523) [subsumption resolution 140668,44565]
140670. ~v2_cat_1(sK819) | ~l1_cat_1(sK818) | ~v2_cat_1(sK817) <- (~1523) [subsumption resolution 140669,44566]
140671. ~l1_cat_1(sK818) | ~v2_cat_1(sK817) <- (~1523) [subsumption resolution 140670,44568]
140672. ~v2_cat_1(sK817) <- (~1523) [subsumption resolution 140671,44564]
140673. $false <- (~1523) [subsumption resolution 140672,44569]
140674. 1523 [avatar contradiction clause 140673]
140699. k9_isocat_2(sK818,X0) = k17_cat_2(sK818,X0) | ~l1_cat_1(sK818) | ~v2_cat_1(X0) | ~l1_cat_1(X0) [resolution 48032,44563]
140703. ~v2_cat_1(X0) | ~l1_cat_1(X0) | k9_isocat_2(sK818,X0) = k17_cat_2(sK818,X0) [subsumption resolution 140699,44564]
140722. ~l1_cat_1(sK819) | k9_isocat_2(sK818,sK819) = k17_cat_2(sK818,sK819) [resolution 140703,44568]
140727. k9_isocat_2(sK818,sK819) = k17_cat_2(sK818,sK819) [subsumption resolution 140722,44565]
140777. ~l1_cat_1(sK817) | ~v2_cat_1(sK819) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~v2_cat_1(sK818) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | k6_isocat_1(sK817,k11_cat_2(sK818,sK819),sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),k9_isocat_2(sK818,sK819)) = k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)) | ~l1_cat_1(sK818) | ~v2_cat_1(sK817) | ~l1_cat_1(sK819) <- (1515, 1517) [resolution 55258,140535]
140780. ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~v2_cat_1(sK819) | ~v2_cat_1(sK817) | ~l1_cat_1(sK819) | ~v2_cat_1(sK818) | k6_isocat_1(sK817,k11_cat_2(sK818,sK819),sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),k9_isocat_2(sK818,sK819)) = k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)) | ~l1_cat_1(sK817) | ~l1_cat_1(sK818) <- (1515, 1517) [duplicate literal removal 140777]
140784. ~v2_cat_1(sK818) | ~l1_cat_1(sK817) | ~l1_cat_1(sK818) | ~l1_cat_1(sK819) | k6_isocat_1(sK817,k11_cat_2(sK818,sK819),sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),k9_isocat_2(sK818,sK819)) = k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)) | ~v2_cat_1(sK817) | ~v2_cat_1(sK819) <- (1515, 1517) [subsumption resolution 140780,44566]
140788. ~v2_cat_1(sK817) | ~v2_cat_1(sK819) | ~l1_cat_1(sK817) | ~l1_cat_1(sK818) | ~l1_cat_1(sK819) | k6_isocat_1(sK817,k11_cat_2(sK818,sK819),sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),k9_isocat_2(sK818,sK819)) = k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)) <- (1515, 1517) [subsumption resolution 140784,44563]
140791. ~l1_cat_1(sK819) | ~l1_cat_1(sK817) | k6_isocat_1(sK817,k11_cat_2(sK818,sK819),sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),k9_isocat_2(sK818,sK819)) = k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)) | ~v2_cat_1(sK819) | ~l1_cat_1(sK818) <- (1515, 1517) [subsumption resolution 140788,44569]
140794. ~l1_cat_1(sK817) | k6_isocat_1(sK817,k11_cat_2(sK818,sK819),sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),k9_isocat_2(sK818,sK819)) = k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)) | ~l1_cat_1(sK818) | ~v2_cat_1(sK819) <- (1515, 1517) [subsumption resolution 140791,44565]
140796. k6_isocat_1(sK817,k11_cat_2(sK818,sK819),sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),k9_isocat_2(sK818,sK819)) = k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)) | ~v2_cat_1(sK819) | ~l1_cat_1(sK818) <- (1515, 1517) [subsumption resolution 140794,44562]
140797. ~l1_cat_1(sK818) | k6_isocat_1(sK817,k11_cat_2(sK818,sK819),sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),k9_isocat_2(sK818,sK819)) = k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)) <- (1515, 1517) [subsumption resolution 140796,44568]
140798. k6_isocat_1(sK817,k11_cat_2(sK818,sK819),sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),k9_isocat_2(sK818,sK819)) = k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)) <- (1515, 1517) [subsumption resolution 140797,44564]
140799. k6_isocat_1(sK817,k11_cat_2(sK818,sK819),sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),k17_cat_2(sK818,sK819)) = k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)) <- (1515, 1517) [forward demodulation 140798,140727]
140810. ~m2_cat_1(X2,X0,X1) | ~l1_cat_1(X1) | ~v2_cat_1(X0) | ~l1_cat_1(X1) | ~m2_cat_1(X2,X0,X1) | ~v2_cat_1(X1) | v1_funct_1(k7_nattra_1(X0,X1,X2)) | ~l1_cat_1(X0) | ~l1_cat_1(X0) | ~v2_cat_1(X1) | ~v2_cat_1(X0) | ~m2_cat_1(X2,X0,X1) [resolution 140473,52012]
140814. v1_funct_1(k7_nattra_1(X0,X1,X2)) | ~l1_cat_1(X0) | ~m2_cat_1(X2,X0,X1) | ~l1_cat_1(X1) | ~v2_cat_1(X0) | ~v2_cat_1(X1) [duplicate literal removal 140810]
140838. ~v2_cat_1(sK817) | ~l1_cat_1(sK817) | ~l1_cat_1(sK818) | v1_funct_1(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820))) | ~v2_cat_1(sK818) | ~m2_cat_1(k11_isocat_2(sK817,sK818,sK819,sK820),sK817,sK818) [superposition 140814,140560]
140839. ~v2_cat_1(sK819) | ~l1_cat_1(sK819) | v1_funct_1(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820))) | ~v2_cat_1(sK817) | ~m2_cat_1(k12_isocat_2(sK817,sK818,sK819,sK820),sK817,sK819) | ~l1_cat_1(sK817) [superposition 140814,140596]
140840. ~v2_cat_1(sK817) | ~m2_cat_1(k12_isocat_2(sK817,sK818,sK819,sK820),sK817,sK819) | ~l1_cat_1(sK819) | ~l1_cat_1(sK817) | v1_funct_1(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820))) [subsumption resolution 140839,44568]
140841. ~v2_cat_1(sK818) | v1_funct_1(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820))) | ~m2_cat_1(k11_isocat_2(sK817,sK818,sK819,sK820),sK817,sK818) | ~l1_cat_1(sK817) | ~l1_cat_1(sK818) [subsumption resolution 140838,44569]
140843. ~l1_cat_1(sK819) | v1_funct_1(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820))) | ~l1_cat_1(sK817) | ~m2_cat_1(k12_isocat_2(sK817,sK818,sK819,sK820),sK817,sK819) [subsumption resolution 140840,44569]
140844. ~l1_cat_1(sK817) | v1_funct_1(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820))) | ~m2_cat_1(k11_isocat_2(sK817,sK818,sK819,sK820),sK817,sK818) | ~l1_cat_1(sK818) [subsumption resolution 140841,44563]
140846. ~m2_cat_1(k12_isocat_2(sK817,sK818,sK819,sK820),sK817,sK819) | v1_funct_1(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820))) | ~l1_cat_1(sK817) [subsumption resolution 140843,44565]
140847. v1_funct_1(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820))) | ~l1_cat_1(sK818) | ~m2_cat_1(k11_isocat_2(sK817,sK818,sK819,sK820),sK817,sK818) [subsumption resolution 140844,44562]
140849. ~l1_cat_1(sK817) | v1_funct_1(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820))) <- (1523) [subsumption resolution 140846,140608]
140850. v1_funct_1(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820))) | ~m2_cat_1(k11_isocat_2(sK817,sK818,sK819,sK820),sK817,sK818) [subsumption resolution 140847,44564]
140852. v1_funct_1(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820))) <- (1523) [subsumption resolution 140849,44562]
140853. v1_funct_1(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820))) <- (1521) [subsumption resolution 140850,140573]
140859. ~l1_cat_1(k11_cat_2(sK818,sK819)) | ~m2_cat_1(X1,k11_cat_2(sK818,sK819),X0) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | r4_nattra_1(u1_cat_1(sK817),u2_cat_1(X0),u1_cat_1(sK817),u2_cat_1(X0),k6_isocat_1(sK817,k11_cat_2(sK818,sK819),X0,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),X1),k7_nattra_1(sK817,X0,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),X0,sK820,X1))) | ~v2_cat_1(k11_cat_2(sK818,sK819)) | ~v2_cat_1(sK817) | ~l1_cat_1(X0) | ~l1_cat_1(sK817) | ~v2_cat_1(X0) <- (1515, 1517) [superposition 46878,140528]
140868. ~l1_cat_1(X0) | ~v2_cat_1(sK817) | ~l1_cat_1(sK817) | ~v2_cat_1(X0) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~v2_cat_1(k11_cat_2(sK818,sK819)) | r4_nattra_1(u1_cat_1(sK817),u2_cat_1(X0),u1_cat_1(sK817),u2_cat_1(X0),k6_isocat_1(sK817,k11_cat_2(sK818,sK819),X0,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),X1),k7_nattra_1(sK817,X0,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),X0,sK820,X1))) | ~m2_cat_1(X1,k11_cat_2(sK818,sK819),X0) <- (1515, 1517) [subsumption resolution 140859,140452]
140874. ~l1_cat_1(X0) | ~v2_cat_1(X0) | ~l1_cat_1(sK817) | r4_nattra_1(u1_cat_1(sK817),u2_cat_1(X0),u1_cat_1(sK817),u2_cat_1(X0),k6_isocat_1(sK817,k11_cat_2(sK818,sK819),X0,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),X1),k7_nattra_1(sK817,X0,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),X0,sK820,X1))) | ~v2_cat_1(k11_cat_2(sK818,sK819)) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~m2_cat_1(X1,k11_cat_2(sK818,sK819),X0) <- (1515, 1517) [subsumption resolution 140868,44569]
140878. ~m2_cat_1(X1,k11_cat_2(sK818,sK819),X0) | ~v2_cat_1(k11_cat_2(sK818,sK819)) | ~v2_cat_1(X0) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~l1_cat_1(X0) | r4_nattra_1(u1_cat_1(sK817),u2_cat_1(X0),u1_cat_1(sK817),u2_cat_1(X0),k6_isocat_1(sK817,k11_cat_2(sK818,sK819),X0,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),X1),k7_nattra_1(sK817,X0,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),X0,sK820,X1))) <- (1515, 1517) [subsumption resolution 140874,44562]
140881. ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~m2_cat_1(X1,k11_cat_2(sK818,sK819),X0) | ~v2_cat_1(X0) | ~l1_cat_1(X0) | r4_nattra_1(u1_cat_1(sK817),u2_cat_1(X0),u1_cat_1(sK817),u2_cat_1(X0),k6_isocat_1(sK817,k11_cat_2(sK818,sK819),X0,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),X1),k7_nattra_1(sK817,X0,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),X0,sK820,X1))) <- (1515, 1517) [subsumption resolution 140878,140444]
140884. r4_nattra_1(u1_cat_1(sK817),u2_cat_1(X0),u1_cat_1(sK817),u2_cat_1(X0),k6_isocat_1(sK817,k11_cat_2(sK818,sK819),X0,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),X1),k7_nattra_1(sK817,X0,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),X0,sK820,X1))) | ~v2_cat_1(X0) | ~m2_cat_1(X1,k11_cat_2(sK818,sK819),X0) | ~l1_cat_1(X0) <- (1515, 1517) [subsumption resolution 140881,44566]
140891. ~m2_cat_1(k8_isocat_2(sK818,sK819),k11_cat_2(sK818,sK819),sK818) | r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k8_isocat_2(sK818,sK819)))) | ~l1_cat_1(sK818) | ~v2_cat_1(sK818) <- (1515, 1517) [superposition 140884,140644]
140892. r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK819),u1_cat_1(sK817),u2_cat_1(sK819),k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),k7_nattra_1(sK817,sK819,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK819,sK820,k17_cat_2(sK818,sK819)))) | ~v2_cat_1(sK819) | ~m2_cat_1(k17_cat_2(sK818,sK819),k11_cat_2(sK818,sK819),sK819) | ~l1_cat_1(sK819) <- (1515, 1517) [superposition 140884,140799]
140900. r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k8_isocat_2(sK818,sK819)))) | ~m2_cat_1(k8_isocat_2(sK818,sK819),k11_cat_2(sK818,sK819),sK818) | ~v2_cat_1(sK818) <- (1515, 1517) [subsumption resolution 140891,44564]
140901. r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK819),u1_cat_1(sK817),u2_cat_1(sK819),k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),k7_nattra_1(sK817,sK819,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK819,sK820,k17_cat_2(sK818,sK819)))) | ~m2_cat_1(k17_cat_2(sK818,sK819),k11_cat_2(sK818,sK819),sK819) | ~l1_cat_1(sK819) <- (1515, 1517) [subsumption resolution 140892,44568]
140903. 1528 <=> v1_xboole_0(u1_cat_1(sK817)) [avatar definition]
140904. ~v1_xboole_0(u1_cat_1(sK817)) <- (~1528) [avatar component clause 140903]
140905. v1_xboole_0(u1_cat_1(sK817)) <- (1528) [avatar component clause 140903]
140918. ~m2_cat_1(k8_isocat_2(sK818,sK819),k11_cat_2(sK818,sK819),sK818) | r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k8_isocat_2(sK818,sK819)))) <- (1515, 1517) [subsumption resolution 140900,44563]
140919. ~m2_cat_1(k17_cat_2(sK818,sK819),k11_cat_2(sK818,sK819),sK819) | r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK819),u1_cat_1(sK817),u2_cat_1(sK819),k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),k7_nattra_1(sK817,sK819,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK819,sK820,k17_cat_2(sK818,sK819)))) <- (1515, 1517) [subsumption resolution 140901,44565]
140921. 1532 <=> m2_cat_1(k8_isocat_2(sK818,sK819),k11_cat_2(sK818,sK819),sK818) [avatar definition]
140922. m2_cat_1(k8_isocat_2(sK818,sK819),k11_cat_2(sK818,sK819),sK818) <- (1532) [avatar component clause 140921]
140923. ~m2_cat_1(k8_isocat_2(sK818,sK819),k11_cat_2(sK818,sK819),sK818) <- (~1532) [avatar component clause 140921]
140925. 1533 <=> r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k8_isocat_2(sK818,sK819)))) [avatar definition]
140927. r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k8_isocat_2(sK818,sK819)))) <- (1533) [avatar component clause 140925]
140928. ~1532 | 1533 | ~1515 | ~1517 [avatar split clause 140918,140451,140443,140925,140921]
140930. 1534 <=> m2_cat_1(k17_cat_2(sK818,sK819),k11_cat_2(sK818,sK819),sK819) [avatar definition]
140932. ~m2_cat_1(k17_cat_2(sK818,sK819),k11_cat_2(sK818,sK819),sK819) <- (~1534) [avatar component clause 140930]
140934. 1535 <=> r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK819),u1_cat_1(sK817),u2_cat_1(sK819),k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),k7_nattra_1(sK817,sK819,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK819,sK820,k17_cat_2(sK818,sK819)))) [avatar definition]
140936. r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK819),u1_cat_1(sK817),u2_cat_1(sK819),k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),k7_nattra_1(sK817,sK819,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK819,sK820,k17_cat_2(sK818,sK819)))) <- (1535) [avatar component clause 140934]
140937. ~1534 | 1535 | ~1515 | ~1517 [avatar split clause 140919,140451,140443,140934,140930]
140938. ~v2_cat_1(sK818) | ~v2_cat_1(sK819) | ~l1_cat_1(sK818) | ~l1_cat_1(sK819) <- (~1532) [resolution 140923,53064]
140939. ~v2_cat_1(sK819) | ~l1_cat_1(sK818) | ~l1_cat_1(sK819) <- (~1532) [subsumption resolution 140938,44563]
140940. ~l1_cat_1(sK818) | ~l1_cat_1(sK819) <- (~1532) [subsumption resolution 140939,44568]
140941. ~l1_cat_1(sK819) <- (~1532) [subsumption resolution 140940,44564]
140942. $false <- (~1532) [subsumption resolution 140941,44565]
140943. 1532 [avatar contradiction clause 140942]
140971. ~l1_cat_1(sK817) <- (1528) [resolution 140905,63394]
140974. $false <- (1528) [subsumption resolution 140971,44562]
140975. ~1528 [avatar contradiction clause 140974]
140977. 1538 <=> v1_xboole_0(u2_cat_1(sK818)) [avatar definition]
140979. v1_xboole_0(u2_cat_1(sK818)) <- (1538) [avatar component clause 140977]
140985. 1540 <=> v1_funct_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) [avatar definition]
140986. v1_funct_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) <- (1540) [avatar component clause 140985]
140987. ~v1_funct_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) <- (~1540) [avatar component clause 140985]
140997. 1543 <=> v1_funct_2(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) [avatar definition]
140999. ~v1_funct_2(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) <- (~1543) [avatar component clause 140997]
141001. 1544 <=> m1_relset_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) [avatar definition]
141003. ~m1_relset_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) <- (~1544) [avatar component clause 141001]
141014. 1547 <=> m2_relset_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) [avatar definition]
141015. m2_relset_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) <- (1547) [avatar component clause 141014]
141030. ~l1_cat_1(X0) | ~v2_cat_1(X0) | k8_isocat_2(sK818,X0) = k16_cat_2(sK818,X0) | ~l1_cat_1(sK818) [resolution 51988,44563]
141036. ~v2_cat_1(X0) | k8_isocat_2(sK818,X0) = k16_cat_2(sK818,X0) | ~l1_cat_1(X0) [subsumption resolution 141030,44564]
141057. k16_cat_2(sK818,sK819) = k8_isocat_2(sK818,sK819) | ~l1_cat_1(sK819) [resolution 141036,44568]
141059. k16_cat_2(sK818,sK819) = k8_isocat_2(sK818,sK819) [subsumption resolution 141057,44565]
141064. r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819)))) <- (1533) [backward demodulation 140927,141059]
141065. m2_cat_1(k16_cat_2(sK818,sK819),k11_cat_2(sK818,sK819),sK818) <- (1532) [backward demodulation 140922,141059]
141066. k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)) = k6_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),k16_cat_2(sK818,sK819)) <- (1515, 1517) [backward demodulation 140644,141059]
141096. r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))),k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | v1_xboole_0(u2_cat_1(sK818)) | ~v1_funct_2(k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))),u1_cat_1(sK817),u2_cat_1(sK818)) | v1_xboole_0(u2_cat_1(sK818)) | ~v1_funct_1(k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819)))) | ~m1_relset_1(k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))),u1_cat_1(sK817),u2_cat_1(sK818)) | v1_xboole_0(u1_cat_1(sK817)) | v1_xboole_0(u1_cat_1(sK817)) | ~m1_relset_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~v1_funct_2(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~v1_funct_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) <- (1533) [resolution 141064,58555]
141098. v1_xboole_0(u1_cat_1(sK817)) | ~v1_funct_2(k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))),u1_cat_1(sK817),u2_cat_1(sK818)) | ~m1_relset_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~v1_funct_1(k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819)))) | ~m1_relset_1(k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))),u1_cat_1(sK817),u2_cat_1(sK818)) | ~v1_funct_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | v1_xboole_0(u2_cat_1(sK818)) | ~v1_funct_2(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))),k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) <- (1533) [duplicate literal removal 141096]
141162. ~v2_cat_1(sK819) | ~l1_cat_1(sK817) | ~v2_cat_1(sK818) | k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK819,sK820,k9_isocat_2(sK818,sK819)) = k12_isocat_2(sK817,sK818,sK819,sK820) | ~l1_cat_1(sK818) | ~l1_cat_1(sK819) | ~v2_cat_1(sK817) [resolution 62120,44566]
141169. ~v2_cat_1(sK818) | ~l1_cat_1(sK818) | k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK819,sK820,k9_isocat_2(sK818,sK819)) = k12_isocat_2(sK817,sK818,sK819,sK820) | ~l1_cat_1(sK817) | ~l1_cat_1(sK819) | ~v2_cat_1(sK817) [subsumption resolution 141162,44568]
141174. ~l1_cat_1(sK819) | ~l1_cat_1(sK817) | k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK819,sK820,k9_isocat_2(sK818,sK819)) = k12_isocat_2(sK817,sK818,sK819,sK820) | ~v2_cat_1(sK817) | ~l1_cat_1(sK818) [subsumption resolution 141169,44563]
141175. ~l1_cat_1(sK817) | k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK819,sK820,k9_isocat_2(sK818,sK819)) = k12_isocat_2(sK817,sK818,sK819,sK820) | ~l1_cat_1(sK818) | ~v2_cat_1(sK817) [subsumption resolution 141174,44565]
141176. ~v2_cat_1(sK817) | ~l1_cat_1(sK818) | k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK819,sK820,k9_isocat_2(sK818,sK819)) = k12_isocat_2(sK817,sK818,sK819,sK820) [subsumption resolution 141175,44562]
141177. ~l1_cat_1(sK818) | k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK819,sK820,k9_isocat_2(sK818,sK819)) = k12_isocat_2(sK817,sK818,sK819,sK820) [subsumption resolution 141176,44569]
141178. k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK819,sK820,k9_isocat_2(sK818,sK819)) = k12_isocat_2(sK817,sK818,sK819,sK820) [subsumption resolution 141177,44564]
141179. k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK819,sK820,k17_cat_2(sK818,sK819)) = k12_isocat_2(sK817,sK818,sK819,sK820) [forward demodulation 141178,140727]
141196. ~v2_cat_1(sK818) | ~l1_cat_1(sK817) | k11_isocat_2(sK817,sK818,sK819,sK820) = k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k8_isocat_2(sK818,sK819)) | ~l1_cat_1(sK819) | ~l1_cat_1(sK818) | ~v2_cat_1(sK817) | ~v2_cat_1(sK819) [resolution 59071,44566]
141201. k11_isocat_2(sK817,sK818,sK819,sK820) = k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k8_isocat_2(sK818,sK819)) | ~l1_cat_1(sK819) | ~v2_cat_1(sK819) | ~l1_cat_1(sK817) | ~v2_cat_1(sK817) | ~l1_cat_1(sK818) [subsumption resolution 141196,44563]
141206. ~l1_cat_1(sK817) | ~v2_cat_1(sK817) | ~v2_cat_1(sK819) | k11_isocat_2(sK817,sK818,sK819,sK820) = k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k8_isocat_2(sK818,sK819)) | ~l1_cat_1(sK818) [subsumption resolution 141201,44565]
141210. k11_isocat_2(sK817,sK818,sK819,sK820) = k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k8_isocat_2(sK818,sK819)) | ~v2_cat_1(sK819) | ~l1_cat_1(sK818) | ~v2_cat_1(sK817) [subsumption resolution 141206,44562]
141212. ~l1_cat_1(sK818) | k11_isocat_2(sK817,sK818,sK819,sK820) = k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k8_isocat_2(sK818,sK819)) | ~v2_cat_1(sK817) [subsumption resolution 141210,44568]
141213. ~v2_cat_1(sK817) | k11_isocat_2(sK817,sK818,sK819,sK820) = k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k8_isocat_2(sK818,sK819)) [subsumption resolution 141212,44564]
141214. k11_isocat_2(sK817,sK818,sK819,sK820) = k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k8_isocat_2(sK818,sK819)) [subsumption resolution 141213,44569]
141215. k11_isocat_2(sK817,sK818,sK819,sK820) = k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819)) [forward demodulation 141214,141059]
141281. ~v2_cat_1(X1) | ~m2_cat_1(k12_isocat_2(X1,X3,X0,X2),X1,X0) | ~v2_cat_1(X3) | ~m2_cat_1(X2,X1,k11_cat_2(X3,X0)) | ~v2_cat_1(X1) | ~m2_nattra_1(X5,X1,k11_cat_2(X3,X0),X4,X2) | ~v2_cat_1(X0) | ~v2_cat_1(X0) | ~m2_cat_1(X4,X1,k11_cat_2(X3,X0)) | ~m2_cat_1(k12_isocat_2(X1,X3,X0,X4),X1,X0) | ~l1_cat_1(X1) | ~l1_cat_1(X3) | ~l1_cat_1(X0) | v1_funct_1(k14_isocat_2(X1,X3,X0,X4,X2,X5)) | ~l1_cat_1(X1) | ~l1_cat_1(X0) [resolution 140503,52012]
141282. ~l1_cat_1(X1) | ~v2_cat_1(X0) | ~m2_nattra_1(X5,X1,k11_cat_2(X3,X0),X4,X2) | ~l1_cat_1(X3) | ~m2_cat_1(k12_isocat_2(X1,X3,X0,X2),X1,X0) | ~v2_cat_1(X1) | ~v2_cat_1(X3) | ~m2_cat_1(X2,X1,k11_cat_2(X3,X0)) | v1_funct_1(k14_isocat_2(X1,X3,X0,X4,X2,X5)) | ~m2_cat_1(k12_isocat_2(X1,X3,X0,X4),X1,X0) | ~m2_cat_1(X4,X1,k11_cat_2(X3,X0)) | ~l1_cat_1(X0) [duplicate literal removal 141281]
141283. ~m2_cat_1(k12_isocat_2(X1,X3,X0,X2),X1,X0) | ~l1_cat_1(X1) | ~l1_cat_1(X3) | ~v2_cat_1(X0) | ~l1_cat_1(X0) | v1_funct_1(k14_isocat_2(X1,X3,X0,X4,X2,X5)) | ~m2_nattra_1(X5,X1,k11_cat_2(X3,X0),X4,X2) | ~m2_cat_1(X2,X1,k11_cat_2(X3,X0)) | ~v2_cat_1(X3) | ~v2_cat_1(X1) | ~m2_cat_1(X4,X1,k11_cat_2(X3,X0)) [subsumption resolution 141282,54428]
141284. ~m2_nattra_1(X5,X1,k11_cat_2(X3,X0),X4,X2) | ~l1_cat_1(X1) | ~v2_cat_1(X3) | ~l1_cat_1(X0) | ~l1_cat_1(X3) | v1_funct_1(k14_isocat_2(X1,X3,X0,X4,X2,X5)) | ~v2_cat_1(X0) | ~m2_cat_1(X2,X1,k11_cat_2(X3,X0)) | ~v2_cat_1(X1) | ~m2_cat_1(X4,X1,k11_cat_2(X3,X0)) [subsumption resolution 141283,54428]
141290. ~l1_cat_1(sK819) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~l1_cat_1(sK817) | ~l1_cat_1(sK818) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | v1_funct_1(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | ~v2_cat_1(sK819) | ~v2_cat_1(sK818) | ~v2_cat_1(sK817) <- (1515, 1517) [resolution 141284,140535]
141294. ~l1_cat_1(sK817) | ~v2_cat_1(sK817) | v1_funct_1(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~l1_cat_1(sK818) | ~v2_cat_1(sK819) | ~v2_cat_1(sK818) | ~l1_cat_1(sK819) <- (1515, 1517) [duplicate literal removal 141290]
141298. ~v2_cat_1(sK819) | ~v2_cat_1(sK818) | ~v2_cat_1(sK817) | ~l1_cat_1(sK818) | ~l1_cat_1(sK819) | v1_funct_1(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) <- (1515, 1517) [subsumption resolution 141294,44562]
141302. v1_funct_1(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | ~v2_cat_1(sK818) | ~l1_cat_1(sK819) | ~l1_cat_1(sK818) | ~v2_cat_1(sK817) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) <- (1515, 1517) [subsumption resolution 141298,44568]
141305. ~l1_cat_1(sK819) | ~v2_cat_1(sK817) | v1_funct_1(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | ~l1_cat_1(sK818) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) <- (1515, 1517) [subsumption resolution 141302,44563]
141308. v1_funct_1(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | ~l1_cat_1(sK818) | ~v2_cat_1(sK817) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) <- (1515, 1517) [subsumption resolution 141305,44565]
141309. ~v2_cat_1(sK817) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | v1_funct_1(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) <- (1515, 1517) [subsumption resolution 141308,44564]
141310. v1_funct_1(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) <- (1515, 1517) [subsumption resolution 141309,44569]
141311. v1_funct_1(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) <- (1515, 1517) [subsumption resolution 141310,44566]
141326. ~m2_cat_1(k11_isocat_2(X0,X2,X1,X3),X0,X2) | v1_funct_1(k13_isocat_2(X0,X2,X1,X5,X3,X4)) | ~l1_cat_1(X2) | ~l1_cat_1(X2) | ~v2_cat_1(X0) | ~v2_cat_1(X2) | ~m2_cat_1(X3,X0,k11_cat_2(X2,X1)) | ~l1_cat_1(X1) | ~v2_cat_1(X2) | ~m2_nattra_1(X4,X0,k11_cat_2(X2,X1),X5,X3) | ~m2_cat_1(X5,X0,k11_cat_2(X2,X1)) | ~l1_cat_1(X0) | ~v2_cat_1(X1) | ~l1_cat_1(X0) | ~v2_cat_1(X0) | ~m2_cat_1(k11_isocat_2(X0,X2,X1,X5),X0,X2) [resolution 140499,52012]
141327. ~l1_cat_1(X1) | ~m2_cat_1(k11_isocat_2(X0,X2,X1,X3),X0,X2) | ~v2_cat_1(X1) | ~m2_cat_1(k11_isocat_2(X0,X2,X1,X5),X0,X2) | ~v2_cat_1(X2) | ~v2_cat_1(X0) | ~m2_cat_1(X5,X0,k11_cat_2(X2,X1)) | ~m2_nattra_1(X4,X0,k11_cat_2(X2,X1),X5,X3) | v1_funct_1(k13_isocat_2(X0,X2,X1,X5,X3,X4)) | ~l1_cat_1(X0) | ~m2_cat_1(X3,X0,k11_cat_2(X2,X1)) | ~l1_cat_1(X2) [duplicate literal removal 141326]
141328. ~m2_nattra_1(X4,X0,k11_cat_2(X2,X1),X5,X3) | ~v2_cat_1(X2) | ~l1_cat_1(X1) | ~l1_cat_1(X2) | ~l1_cat_1(X0) | ~m2_cat_1(X5,X0,k11_cat_2(X2,X1)) | v1_funct_1(k13_isocat_2(X0,X2,X1,X5,X3,X4)) | ~m2_cat_1(k11_isocat_2(X0,X2,X1,X5),X0,X2) | ~m2_cat_1(X3,X0,k11_cat_2(X2,X1)) | ~v2_cat_1(X0) | ~v2_cat_1(X1) [subsumption resolution 141327,44450]
141329. ~m2_nattra_1(X4,X0,k11_cat_2(X2,X1),X5,X3) | ~v2_cat_1(X2) | ~v2_cat_1(X1) | ~l1_cat_1(X2) | ~l1_cat_1(X0) | ~m2_cat_1(X3,X0,k11_cat_2(X2,X1)) | ~v2_cat_1(X0) | ~l1_cat_1(X1) | v1_funct_1(k13_isocat_2(X0,X2,X1,X5,X3,X4)) | ~m2_cat_1(X5,X0,k11_cat_2(X2,X1)) [subsumption resolution 141328,44450]
141339. ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | v1_funct_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | ~l1_cat_1(sK818) | ~v2_cat_1(sK819) | ~v2_cat_1(sK817) | ~l1_cat_1(sK819) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~v2_cat_1(sK818) | ~l1_cat_1(sK817) <- (1515, 1517) [resolution 141329,140535]
141344. ~l1_cat_1(sK819) | ~v2_cat_1(sK819) | ~v2_cat_1(sK818) | ~l1_cat_1(sK818) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~v2_cat_1(sK817) | v1_funct_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | ~l1_cat_1(sK817) <- (1515, 1517) [duplicate literal removal 141339]
141349. ~l1_cat_1(sK818) | v1_funct_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | ~v2_cat_1(sK819) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~l1_cat_1(sK817) | ~v2_cat_1(sK817) | ~v2_cat_1(sK818) <- (1515, 1517) [subsumption resolution 141344,44565]
141354. ~l1_cat_1(sK817) | ~v2_cat_1(sK817) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~v2_cat_1(sK818) | ~v2_cat_1(sK819) | v1_funct_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) <- (1515, 1517) [subsumption resolution 141349,44564]
141357. ~v2_cat_1(sK819) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~v2_cat_1(sK817) | ~v2_cat_1(sK818) | v1_funct_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) <- (1515, 1517) [subsumption resolution 141354,44562]
141360. ~v2_cat_1(sK817) | v1_funct_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | ~v2_cat_1(sK818) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) <- (1515, 1517) [subsumption resolution 141357,44568]
141361. ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~v2_cat_1(sK818) | ~v2_cat_1(sK817) <- (1515, 1517, ~1540) [subsumption resolution 141360,140987]
141362. ~v2_cat_1(sK818) | ~v2_cat_1(sK817) <- (1515, 1517, ~1540) [subsumption resolution 141361,44566]
141363. ~v2_cat_1(sK817) <- (1515, 1517, ~1540) [subsumption resolution 141362,44563]
141364. $false <- (1515, 1517, ~1540) [subsumption resolution 141363,44569]
141365. ~1515 | ~1517 | 1540 [avatar contradiction clause 141364]
141369. ~m1_relset_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~v1_funct_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | ~v1_funct_2(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~m1_relset_1(k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))),u1_cat_1(sK817),u2_cat_1(sK818)) | ~v1_funct_1(k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819)))) | v1_xboole_0(u2_cat_1(sK818)) | r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))),k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | ~v1_funct_2(k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))),u1_cat_1(sK817),u2_cat_1(sK818)) <- (~1528, 1533) [subsumption resolution 141098,140904]
141376. ~v1_funct_2(k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))),u1_cat_1(sK817),u2_cat_1(sK818)) | ~v1_funct_1(k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819)))) | v1_xboole_0(u2_cat_1(sK818)) | ~m1_relset_1(k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))),u1_cat_1(sK817),u2_cat_1(sK818)) | ~m1_relset_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))),k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | ~v1_funct_2(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) <- (~1528, 1533, 1540) [subsumption resolution 141369,140986]
141383. v1_xboole_0(u2_cat_1(sK818)) | ~v1_funct_2(k7_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~m1_relset_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~m1_relset_1(k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))),u1_cat_1(sK817),u2_cat_1(sK818)) | ~v1_funct_1(k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819)))) | r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))),k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | ~v1_funct_2(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) <- (~1528, 1533, 1540) [forward demodulation 141376,141215]
141391. ~m1_relset_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~v1_funct_2(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | v1_xboole_0(u2_cat_1(sK818)) | ~v1_funct_2(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~v1_funct_1(k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819)))) | ~m1_relset_1(k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))),u1_cat_1(sK817),u2_cat_1(sK818)) | r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))),k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) <- (~1528, 1533, 1540) [forward demodulation 141383,140560]
141394. 1550 <=> v1_funct_2(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) [avatar definition]
141396. ~v1_funct_2(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) <- (~1550) [avatar component clause 141394]
141402. 1552 <=> m1_relset_1(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) [avatar definition]
141410. 1554 <=> m2_relset_1(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) [avatar definition]
141411. m2_relset_1(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) <- (1554) [avatar component clause 141410]
141416. ~m1_relset_1(k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))),u1_cat_1(sK817),u2_cat_1(sK818)) | ~m1_relset_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))),k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | ~v1_funct_1(k7_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820))) | ~v1_funct_2(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~v1_funct_2(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | v1_xboole_0(u2_cat_1(sK818)) <- (~1528, 1533, 1540) [forward demodulation 141391,141215]
141420. v1_xboole_0(u2_cat_1(sK818)) | ~v1_funct_2(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~m1_relset_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~v1_funct_1(k7_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820))) | ~v1_funct_2(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~m1_relset_1(k7_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))),k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) <- (~1528, 1533, 1540) [forward demodulation 141416,141215]
141423. r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))),k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | ~m1_relset_1(k7_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~m1_relset_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~v1_funct_2(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | v1_xboole_0(u2_cat_1(sK818)) | ~v1_funct_2(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~v1_funct_1(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820))) <- (~1528, 1533, 1540) [forward demodulation 141420,140560]
141426. ~v1_funct_2(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~v1_funct_2(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~m1_relset_1(k7_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~m1_relset_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | v1_xboole_0(u2_cat_1(sK818)) | r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))),k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) <- (1521, ~1528, 1533, 1540) [subsumption resolution 141423,140853]
141429. v1_xboole_0(u2_cat_1(sK818)) | ~m1_relset_1(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~m1_relset_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~v1_funct_2(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k7_nattra_1(sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))),k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | ~v1_funct_2(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) <- (1521, ~1528, 1533, 1540) [forward demodulation 141426,140560]
141432. v1_xboole_0(u2_cat_1(sK818)) | r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k7_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | ~m1_relset_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~v1_funct_2(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~v1_funct_2(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~m1_relset_1(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) <- (1521, ~1528, 1533, 1540) [forward demodulation 141429,141215]
141435. ~v1_funct_2(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | v1_xboole_0(u2_cat_1(sK818)) | ~m1_relset_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~m1_relset_1(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~v1_funct_2(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK818),u1_cat_1(sK817),u2_cat_1(sK818),k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) <- (1521, ~1528, 1533, 1540) [forward demodulation 141432,140560]
141438. ~v1_funct_2(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~m1_relset_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~v1_funct_2(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | v1_xboole_0(u2_cat_1(sK818)) | ~m1_relset_1(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) <- (~370, 1515, 1517, 1521, ~1528, 1533, 1540) [subsumption resolution 141435,140561]
141440. ~1552 | ~1550 | 1538 | ~1544 | ~1543 | 370 | ~1515 | ~1517 | ~1521 | 1528 | ~1533 | ~1540 [avatar split clause 141438,140985,140925,140903,140572,140451,140443,78356,140997,141001,140977,141394,141402]
142311. ~v2_cat_1(sK819) | ~v2_cat_1(sK818) | ~l1_cat_1(sK818) | ~l1_cat_1(sK819) <- (~1534) [resolution 44441,140932]
142323. ~l1_cat_1(sK818) | ~v2_cat_1(sK818) | ~l1_cat_1(sK819) <- (~1534) [subsumption resolution 142311,44568]
142332. ~v2_cat_1(sK818) | ~l1_cat_1(sK819) <- (~1534) [subsumption resolution 142323,44564]
142341. ~l1_cat_1(sK819) <- (~1534) [subsumption resolution 142332,44563]
142345. $false <- (~1534) [subsumption resolution 142341,44565]
142346. 1534 [avatar contradiction clause 142345]
142352. r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK819),u1_cat_1(sK817),u2_cat_1(sK819),k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),k7_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820))) <- (1535) [forward demodulation 140936,141179]
142358. r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK819),u1_cat_1(sK817),u2_cat_1(sK819),k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820))) <- (1535) [forward demodulation 142352,140596]
142385. r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK819),u1_cat_1(sK817),u2_cat_1(sK819),k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | v1_xboole_0(u1_cat_1(sK817)) | ~v1_funct_2(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) | ~m1_relset_1(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) | v1_xboole_0(u2_cat_1(sK819)) | v1_xboole_0(u1_cat_1(sK817)) | ~m1_relset_1(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) | ~v1_funct_1(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820))) | ~v1_funct_1(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | v1_xboole_0(u2_cat_1(sK819)) | ~v1_funct_2(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) <- (1535) [resolution 142358,58555]
142387. v1_xboole_0(u1_cat_1(sK817)) | r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK819),u1_cat_1(sK817),u2_cat_1(sK819),k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | ~v1_funct_2(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) | v1_xboole_0(u2_cat_1(sK819)) | ~m1_relset_1(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) | ~m1_relset_1(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) | ~v1_funct_1(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820))) | ~v1_funct_2(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) | ~v1_funct_1(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) <- (1535) [duplicate literal removal 142385]
142390. ~v1_funct_2(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) | ~v1_funct_1(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | ~m1_relset_1(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) | v1_xboole_0(u2_cat_1(sK819)) | ~v1_funct_2(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) | r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK819),u1_cat_1(sK817),u2_cat_1(sK819),k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | ~m1_relset_1(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) | ~v1_funct_1(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820))) <- (~1528, 1535) [subsumption resolution 142387,140904]
142393. v1_xboole_0(u2_cat_1(sK819)) | ~v1_funct_2(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) | ~m1_relset_1(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) | ~v1_funct_1(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820))) | ~v1_funct_2(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) | ~m1_relset_1(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) | r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK819),u1_cat_1(sK817),u2_cat_1(sK819),k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) <- (1515, 1517, ~1528, 1535) [subsumption resolution 142390,141311]
142396. ~v1_funct_2(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) | ~v1_funct_2(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) | v1_xboole_0(u2_cat_1(sK819)) | ~m1_relset_1(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) | r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK819),u1_cat_1(sK817),u2_cat_1(sK819),k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) | ~m1_relset_1(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) <- (1515, 1517, 1523, ~1528, 1535) [subsumption resolution 142393,140852]
142400. 1618 <=> v1_funct_2(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) [avatar definition]
142402. ~v1_funct_2(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) <- (~1618) [avatar component clause 142400]
142404. 1619 <=> m1_relset_1(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) [avatar definition]
142408. 1620 <=> v1_funct_2(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) [avatar definition]
142410. ~v1_funct_2(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) <- (~1620) [avatar component clause 142408]
142412. 1621 <=> v1_xboole_0(u2_cat_1(sK819)) [avatar definition]
142414. v1_xboole_0(u2_cat_1(sK819)) <- (1621) [avatar component clause 142412]
142416. 1622 <=> m1_relset_1(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) [avatar definition]
142418. ~m1_relset_1(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) <- (~1622) [avatar component clause 142416]
142420. 1623 <=> r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK819),u1_cat_1(sK817),u2_cat_1(sK819),k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) [avatar definition]
142423. ~1618 | ~1619 | ~1620 | 1621 | ~1622 | 1623 | ~1515 | ~1517 | ~1523 | 1528 | ~1535 [avatar split clause 142396,140934,140903,140607,140451,140443,142420,142416,142412,142408,142404,142400]
142425. 1624 <=> m2_relset_1(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) [avatar definition]
142426. m2_relset_1(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) <- (1624) [avatar component clause 142425]
142429. 1625 <=> m2_relset_1(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) [avatar definition]
142430. m2_relset_1(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) <- (1625) [avatar component clause 142429]
142431. ~m2_relset_1(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) <- (~1625) [avatar component clause 142429]
144146. ~l1_cat_1(X2) | ~m2_cat_1(X0,X1,X2) | ~v2_cat_1(X2) | ~l1_cat_1(X2) | ~l1_cat_1(X1) | ~v2_cat_1(X1) | m2_relset_1(k7_nattra_1(X1,X2,X0),u1_cat_1(X1),u2_cat_1(X2)) | ~v2_cat_1(X2) | ~m2_cat_1(X0,X1,X2) | ~m2_cat_1(X0,X1,X2) | ~l1_cat_1(X1) | ~v2_cat_1(X1) [resolution 52013,140473]
144147. ~l1_cat_1(X0) | ~m2_cat_1(X4,X0,k11_cat_2(X1,X2)) | ~v2_cat_1(X2) | ~m2_cat_1(k12_isocat_2(X0,X1,X2,X3),X0,X2) | m2_relset_1(k14_isocat_2(X0,X1,X2,X4,X3,X5),u1_cat_1(X0),u2_cat_1(X2)) | ~v2_cat_1(X1) | ~v2_cat_1(X2) | ~l1_cat_1(X1) | ~m2_cat_1(k12_isocat_2(X0,X1,X2,X4),X0,X2) | ~l1_cat_1(X0) | ~m2_cat_1(X3,X0,k11_cat_2(X1,X2)) | ~l1_cat_1(X2) | ~v2_cat_1(X0) | ~m2_nattra_1(X5,X0,k11_cat_2(X1,X2),X4,X3) | ~l1_cat_1(X2) | ~v2_cat_1(X0) [resolution 52013,140503]
144149. m2_relset_1(k7_nattra_1(X1,X2,X0),u1_cat_1(X1),u2_cat_1(X2)) | ~v2_cat_1(X2) | ~l1_cat_1(X2) | ~v2_cat_1(X1) | ~l1_cat_1(X1) | ~m2_cat_1(X0,X1,X2) [duplicate literal removal 144146]
144151. ~m2_cat_1(k12_isocat_2(X0,X1,X2,X4),X0,X2) | ~v2_cat_1(X1) | ~m2_cat_1(X4,X0,k11_cat_2(X1,X2)) | ~l1_cat_1(X2) | ~v2_cat_1(X0) | ~l1_cat_1(X1) | ~m2_nattra_1(X5,X0,k11_cat_2(X1,X2),X4,X3) | ~v2_cat_1(X2) | ~m2_cat_1(k12_isocat_2(X0,X1,X2,X3),X0,X2) | ~l1_cat_1(X0) | m2_relset_1(k14_isocat_2(X0,X1,X2,X4,X3,X5),u1_cat_1(X0),u2_cat_1(X2)) | ~m2_cat_1(X3,X0,k11_cat_2(X1,X2)) [duplicate literal removal 144147]
144153. ~l1_cat_1(X2) | ~l1_cat_1(X1) | m2_relset_1(k14_isocat_2(X0,X1,X2,X4,X3,X5),u1_cat_1(X0),u2_cat_1(X2)) | ~m2_cat_1(X3,X0,k11_cat_2(X1,X2)) | ~v2_cat_1(X0) | ~v2_cat_1(X2) | ~l1_cat_1(X0) | ~m2_nattra_1(X5,X0,k11_cat_2(X1,X2),X4,X3) | ~m2_cat_1(k12_isocat_2(X0,X1,X2,X4),X0,X2) | ~m2_cat_1(X4,X0,k11_cat_2(X1,X2)) | ~v2_cat_1(X1) [subsumption resolution 144151,54428]
144155. m2_relset_1(k14_isocat_2(X0,X1,X2,X4,X3,X5),u1_cat_1(X0),u2_cat_1(X2)) | ~l1_cat_1(X0) | ~l1_cat_1(X2) | ~v2_cat_1(X1) | ~m2_cat_1(X3,X0,k11_cat_2(X1,X2)) | ~m2_nattra_1(X5,X0,k11_cat_2(X1,X2),X4,X3) | ~v2_cat_1(X0) | ~v2_cat_1(X2) | ~m2_cat_1(X4,X0,k11_cat_2(X1,X2)) | ~l1_cat_1(X1) [subsumption resolution 144153,54428]
144163. m2_relset_1(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~m2_cat_1(k11_isocat_2(sK817,sK818,sK819,sK820),sK817,sK818) | ~l1_cat_1(sK818) | ~v2_cat_1(sK817) | ~v2_cat_1(sK818) | ~l1_cat_1(sK817) [superposition 144149,140560]
144164. ~v2_cat_1(sK819) | m2_relset_1(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) | ~m2_cat_1(k12_isocat_2(sK817,sK818,sK819,sK820),sK817,sK819) | ~v2_cat_1(sK817) | ~l1_cat_1(sK819) | ~l1_cat_1(sK817) [superposition 144149,140596]
144179. m2_relset_1(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) | ~l1_cat_1(sK817) | ~v2_cat_1(sK817) | ~m2_cat_1(k12_isocat_2(sK817,sK818,sK819,sK820),sK817,sK819) | ~l1_cat_1(sK819) [subsumption resolution 144164,44568]
144181. ~v2_cat_1(sK817) | ~l1_cat_1(sK818) | ~v2_cat_1(sK818) | m2_relset_1(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~l1_cat_1(sK817) <- (1521) [subsumption resolution 144163,140573]
144186. ~l1_cat_1(sK819) | m2_relset_1(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) | ~v2_cat_1(sK817) | ~m2_cat_1(k12_isocat_2(sK817,sK818,sK819,sK820),sK817,sK819) [subsumption resolution 144179,44562]
144188. m2_relset_1(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~l1_cat_1(sK818) | ~l1_cat_1(sK817) | ~v2_cat_1(sK818) <- (1521) [subsumption resolution 144181,44569]
144193. ~v2_cat_1(sK817) | m2_relset_1(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) | ~m2_cat_1(k12_isocat_2(sK817,sK818,sK819,sK820),sK817,sK819) [subsumption resolution 144186,44565]
144195. m2_relset_1(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~l1_cat_1(sK817) | ~v2_cat_1(sK818) <- (1521) [subsumption resolution 144188,44564]
144200. m2_relset_1(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) | ~m2_cat_1(k12_isocat_2(sK817,sK818,sK819,sK820),sK817,sK819) [subsumption resolution 144193,44569]
144202. ~v2_cat_1(sK818) | m2_relset_1(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) <- (1521) [subsumption resolution 144195,44562]
144207. m2_relset_1(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) <- (1523) [subsumption resolution 144200,140608]
144209. m2_relset_1(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) <- (1521) [subsumption resolution 144202,44563]
144214. 1624 | ~1523 [avatar split clause 144207,140607,142425]
144215. 1554 | ~1521 [avatar split clause 144209,140572,141410]
144221. m1_relset_1(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) <- (1624) [resolution 142426,65267]
144222. 1619 | ~1624 [avatar split clause 144221,142425,142404]
144235. m1_relset_1(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) <- (1554) [resolution 141411,65267]
144236. 1552 | ~1554 [avatar split clause 144235,141410,141402]
146685. ~m2_cat_1(X1,X0,X2) | ~v2_cat_1(X0) | ~l1_cat_1(X2) | ~v2_cat_1(X0) | v1_funct_2(k7_nattra_1(X0,X2,X1),u1_cat_1(X0),u2_cat_1(X2)) | ~m2_cat_1(X1,X0,X2) | ~v2_cat_1(X2) | ~l1_cat_1(X0) | ~l1_cat_1(X0) | ~l1_cat_1(X2) | ~v2_cat_1(X2) | ~m2_cat_1(X1,X0,X2) [resolution 52011,140473]
146686. ~l1_cat_1(X0) | v1_funct_2(k14_isocat_2(X0,X1,X2,X4,X3,X5),u1_cat_1(X0),u2_cat_1(X2)) | ~m2_cat_1(X3,X0,k11_cat_2(X1,X2)) | ~v2_cat_1(X0) | ~l1_cat_1(X2) | ~v2_cat_1(X2) | ~m2_nattra_1(X5,X0,k11_cat_2(X1,X2),X4,X3) | ~l1_cat_1(X1) | ~m2_cat_1(k12_isocat_2(X0,X1,X2,X4),X0,X2) | ~l1_cat_1(X0) | ~m2_cat_1(X4,X0,k11_cat_2(X1,X2)) | ~v2_cat_1(X2) | ~m2_cat_1(k12_isocat_2(X0,X1,X2,X3),X0,X2) | ~l1_cat_1(X2) | ~v2_cat_1(X1) | ~v2_cat_1(X0) [resolution 52011,140503]
146690. ~l1_cat_1(X1) | ~m2_cat_1(X4,X0,k11_cat_2(X1,X2)) | ~l1_cat_1(X2) | ~m2_nattra_1(X5,X0,k11_cat_2(X1,X2),X4,X3) | ~l1_cat_1(X0) | ~v2_cat_1(X1) | ~v2_cat_1(X2) | ~v2_cat_1(X0) | ~m2_cat_1(k12_isocat_2(X0,X1,X2,X4),X0,X2) | ~m2_cat_1(k12_isocat_2(X0,X1,X2,X3),X0,X2) | ~m2_cat_1(X3,X0,k11_cat_2(X1,X2)) | v1_funct_2(k14_isocat_2(X0,X1,X2,X4,X3,X5),u1_cat_1(X0),u2_cat_1(X2)) [duplicate literal removal 146686]
146691. v1_funct_2(k7_nattra_1(X0,X2,X1),u1_cat_1(X0),u2_cat_1(X2)) | ~v2_cat_1(X0) | ~m2_cat_1(X1,X0,X2) | ~v2_cat_1(X2) | ~l1_cat_1(X2) | ~l1_cat_1(X0) [duplicate literal removal 146685]
146694. ~m2_cat_1(k12_isocat_2(X0,X1,X2,X3),X0,X2) | ~v2_cat_1(X2) | v1_funct_2(k14_isocat_2(X0,X1,X2,X4,X3,X5),u1_cat_1(X0),u2_cat_1(X2)) | ~v2_cat_1(X0) | ~l1_cat_1(X1) | ~v2_cat_1(X1) | ~m2_nattra_1(X5,X0,k11_cat_2(X1,X2),X4,X3) | ~m2_cat_1(X4,X0,k11_cat_2(X1,X2)) | ~l1_cat_1(X2) | ~l1_cat_1(X0) | ~m2_cat_1(X3,X0,k11_cat_2(X1,X2)) [subsumption resolution 146690,54428]
146696. v1_funct_2(k14_isocat_2(X0,X1,X2,X4,X3,X5),u1_cat_1(X0),u2_cat_1(X2)) | ~l1_cat_1(X2) | ~m2_cat_1(X3,X0,k11_cat_2(X1,X2)) | ~v2_cat_1(X2) | ~l1_cat_1(X1) | ~v2_cat_1(X0) | ~m2_cat_1(X4,X0,k11_cat_2(X1,X2)) | ~v2_cat_1(X1) | ~l1_cat_1(X0) | ~m2_nattra_1(X5,X0,k11_cat_2(X1,X2),X4,X3) [subsumption resolution 146694,54428]
146699. ~m2_cat_1(k11_isocat_2(sK817,sK818,sK819,sK820),sK817,sK818) | v1_funct_2(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~l1_cat_1(sK817) | ~l1_cat_1(sK818) | ~v2_cat_1(sK818) | ~v2_cat_1(sK817) [superposition 146691,140560]
146700. ~l1_cat_1(sK817) | ~l1_cat_1(sK819) | ~v2_cat_1(sK819) | v1_funct_2(k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) | ~v2_cat_1(sK817) | ~m2_cat_1(k12_isocat_2(sK817,sK818,sK819,sK820),sK817,sK819) [superposition 146691,140596]
146726. ~v2_cat_1(sK817) | ~l1_cat_1(sK818) | ~v2_cat_1(sK818) | ~l1_cat_1(sK817) | v1_funct_2(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) <- (1521) [subsumption resolution 146699,140573]
146738. ~l1_cat_1(sK819) | ~v2_cat_1(sK819) | ~l1_cat_1(sK817) | ~m2_cat_1(k12_isocat_2(sK817,sK818,sK819,sK820),sK817,sK819) | ~v2_cat_1(sK817) <- (~1618) [subsumption resolution 146700,142402]
146740. ~l1_cat_1(sK818) | v1_funct_2(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~v2_cat_1(sK818) | ~l1_cat_1(sK817) <- (1521) [subsumption resolution 146726,44569]
146752. ~v2_cat_1(sK817) | ~m2_cat_1(k12_isocat_2(sK817,sK818,sK819,sK820),sK817,sK819) | ~l1_cat_1(sK817) | ~v2_cat_1(sK819) <- (~1618) [subsumption resolution 146738,44565]
146754. ~v2_cat_1(sK818) | v1_funct_2(k4_nattra_1(sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~l1_cat_1(sK817) <- (1521) [subsumption resolution 146740,44564]
146766. ~m2_cat_1(k12_isocat_2(sK817,sK818,sK819,sK820),sK817,sK819) | ~v2_cat_1(sK819) | ~l1_cat_1(sK817) <- (~1618) [subsumption resolution 146752,44569]
146768. ~v2_cat_1(sK818) | ~l1_cat_1(sK817) <- (1521, ~1550) [subsumption resolution 146754,141396]
146780. ~v2_cat_1(sK819) | ~l1_cat_1(sK817) <- (1523, ~1618) [subsumption resolution 146766,140608]
146782. ~l1_cat_1(sK817) <- (1521, ~1550) [subsumption resolution 146768,44563]
146793. ~l1_cat_1(sK817) <- (1523, ~1618) [subsumption resolution 146780,44568]
146795. $false <- (1521, ~1550) [subsumption resolution 146782,44562]
146796. ~1521 | 1550 [avatar contradiction clause 146795]
146797. $false <- (1523, ~1618) [subsumption resolution 146793,44562]
146798. ~1523 | 1618 [avatar contradiction clause 146797]
146851. ~l1_cat_1(sK819) <- (1621) [resolution 142414,48525]
146854. $false <- (1621) [subsumption resolution 146851,44565]
146855. ~1621 [avatar contradiction clause 146854]
146869. ~l1_cat_1(sK818) <- (1538) [resolution 140979,48525]
146872. $false <- (1538) [subsumption resolution 146869,44564]
146873. ~1538 [avatar contradiction clause 146872]
151169. ~l1_cat_1(sK817) | ~m2_cat_1(k16_cat_2(sK818,sK819),k11_cat_2(sK818,sK819),sK818) | ~v2_cat_1(k11_cat_2(sK818,sK819)) | m2_nattra_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819)),k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))) | ~v2_cat_1(sK818) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~v2_cat_1(sK817) | ~m2_nattra_1(k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),sK817,k11_cat_2(sK818,sK819),sK820,sK820) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~l1_cat_1(k11_cat_2(sK818,sK819)) | ~l1_cat_1(sK818) <- (1515, 1517) [superposition 49971,141066]
151178. ~l1_cat_1(sK817) | ~v2_cat_1(sK818) | ~v2_cat_1(k11_cat_2(sK818,sK819)) | ~m2_nattra_1(k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),sK817,k11_cat_2(sK818,sK819),sK820,sK820) | ~v2_cat_1(sK817) | m2_nattra_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819)),k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))) | ~l1_cat_1(k11_cat_2(sK818,sK819)) | ~l1_cat_1(sK818) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~m2_cat_1(k16_cat_2(sK818,sK819),k11_cat_2(sK818,sK819),sK818) <- (1515, 1517) [duplicate literal removal 151169]
151188. ~m2_nattra_1(k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),sK817,k11_cat_2(sK818,sK819),sK820,sK820) | ~m2_cat_1(k16_cat_2(sK818,sK819),k11_cat_2(sK818,sK819),sK818) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~v2_cat_1(sK817) | ~l1_cat_1(k11_cat_2(sK818,sK819)) | ~v2_cat_1(k11_cat_2(sK818,sK819)) | ~v2_cat_1(sK818) | ~l1_cat_1(sK818) | m2_nattra_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819)),k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))) <- (1515, 1517) [subsumption resolution 151178,44562]
151201. ~v2_cat_1(sK818) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~v2_cat_1(k11_cat_2(sK818,sK819)) | ~v2_cat_1(sK817) | m2_nattra_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819)),k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))) | ~l1_cat_1(sK818) | ~l1_cat_1(k11_cat_2(sK818,sK819)) | ~m2_cat_1(k16_cat_2(sK818,sK819),k11_cat_2(sK818,sK819),sK818) <- (1515, 1517) [subsumption resolution 151188,140535]
151213. m2_nattra_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819)),k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))) | ~v2_cat_1(sK817) | ~l1_cat_1(k11_cat_2(sK818,sK819)) | ~m2_cat_1(k16_cat_2(sK818,sK819),k11_cat_2(sK818,sK819),sK818) | ~l1_cat_1(sK818) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~v2_cat_1(k11_cat_2(sK818,sK819)) <- (1515, 1517) [subsumption resolution 151201,44563]
151223. m2_nattra_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819)),k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))) | ~l1_cat_1(sK818) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~l1_cat_1(k11_cat_2(sK818,sK819)) | ~m2_cat_1(k16_cat_2(sK818,sK819),k11_cat_2(sK818,sK819),sK818) | ~v2_cat_1(k11_cat_2(sK818,sK819)) <- (1515, 1517) [subsumption resolution 151213,44569]
151231. ~v2_cat_1(k11_cat_2(sK818,sK819)) | ~m2_cat_1(k16_cat_2(sK818,sK819),k11_cat_2(sK818,sK819),sK818) | ~l1_cat_1(k11_cat_2(sK818,sK819)) | m2_nattra_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819)),k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) <- (1515, 1517) [subsumption resolution 151223,44564]
151237. ~l1_cat_1(k11_cat_2(sK818,sK819)) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~m2_cat_1(k16_cat_2(sK818,sK819),k11_cat_2(sK818,sK819),sK818) | m2_nattra_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819)),k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))) <- (1515, 1517) [subsumption resolution 151231,140444]
151243. ~m2_cat_1(k16_cat_2(sK818,sK819),k11_cat_2(sK818,sK819),sK818) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | m2_nattra_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819)),k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))) <- (1515, 1517) [subsumption resolution 151237,140452]
151249. m2_nattra_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819)),k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) <- (1515, 1517, 1532) [subsumption resolution 151243,141065]
151254. m2_nattra_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),sK817,sK818,k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819)),k2_isocat_1(sK817,k11_cat_2(sK818,sK819),sK818,sK820,k16_cat_2(sK818,sK819))) <- (1515, 1517, 1532) [subsumption resolution 151249,44566]
151256. m2_nattra_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820),k11_isocat_2(sK817,sK818,sK819,sK820)) <- (1515, 1517, 1532) [forward demodulation 151254,141215]
151259. ~l1_cat_1(sK817) | ~m2_cat_1(k11_isocat_2(sK817,sK818,sK819,sK820),sK817,sK818) | ~v2_cat_1(sK817) | ~v2_cat_1(sK818) | ~l1_cat_1(sK818) | m1_nattra_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820),k11_isocat_2(sK817,sK818,sK819,sK820)) | ~m2_cat_1(k11_isocat_2(sK817,sK818,sK819,sK820),sK817,sK818) <- (1515, 1517, 1532) [resolution 151256,62059]
151260. ~v2_cat_1(sK817) | ~m2_cat_1(k11_isocat_2(sK817,sK818,sK819,sK820),sK817,sK818) | ~l1_cat_1(sK818) | ~l1_cat_1(sK817) | m1_nattra_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820),k11_isocat_2(sK817,sK818,sK819,sK820)) | ~v2_cat_1(sK818) <- (1515, 1517, 1532) [duplicate literal removal 151259]
151263. ~l1_cat_1(sK818) | ~m2_cat_1(k11_isocat_2(sK817,sK818,sK819,sK820),sK817,sK818) | ~v2_cat_1(sK818) | ~l1_cat_1(sK817) | m1_nattra_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820),k11_isocat_2(sK817,sK818,sK819,sK820)) <- (1515, 1517, 1532) [subsumption resolution 151260,44569]
151266. ~m2_cat_1(k11_isocat_2(sK817,sK818,sK819,sK820),sK817,sK818) | ~l1_cat_1(sK817) | ~v2_cat_1(sK818) | m1_nattra_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820),k11_isocat_2(sK817,sK818,sK819,sK820)) <- (1515, 1517, 1532) [subsumption resolution 151263,44564]
151269. ~l1_cat_1(sK817) | ~v2_cat_1(sK818) | m1_nattra_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820),k11_isocat_2(sK817,sK818,sK819,sK820)) <- (1515, 1517, 1521, 1532) [subsumption resolution 151266,140573]
151272. ~v2_cat_1(sK818) | m1_nattra_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820),k11_isocat_2(sK817,sK818,sK819,sK820)) <- (1515, 1517, 1521, 1532) [subsumption resolution 151269,44562]
151275. m1_nattra_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),sK817,sK818,k11_isocat_2(sK817,sK818,sK819,sK820),k11_isocat_2(sK817,sK818,sK819,sK820)) <- (1515, 1517, 1521, 1532) [subsumption resolution 151272,44563]
151307. v1_funct_2(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~m2_cat_1(k11_isocat_2(sK817,sK818,sK819,sK820),sK817,sK818) | ~m2_cat_1(k11_isocat_2(sK817,sK818,sK819,sK820),sK817,sK818) | ~l1_cat_1(sK817) | ~l1_cat_1(sK818) | ~v2_cat_1(sK818) | ~v2_cat_1(sK817) <- (1515, 1517, 1521, 1532) [resolution 151275,52011]
151308. ~v2_cat_1(sK817) | ~m2_cat_1(k11_isocat_2(sK817,sK818,sK819,sK820),sK817,sK818) | ~l1_cat_1(sK818) | ~m2_cat_1(k11_isocat_2(sK817,sK818,sK819,sK820),sK817,sK818) | ~v2_cat_1(sK818) | m2_relset_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~l1_cat_1(sK817) <- (1515, 1517, 1521, 1532) [resolution 151275,52013]
151311. ~v2_cat_1(sK818) | ~v2_cat_1(sK817) | ~l1_cat_1(sK817) | v1_funct_2(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~l1_cat_1(sK818) | ~m2_cat_1(k11_isocat_2(sK817,sK818,sK819,sK820),sK817,sK818) <- (1515, 1517, 1521, 1532) [duplicate literal removal 151307]
151313. ~v2_cat_1(sK818) | ~l1_cat_1(sK818) | ~l1_cat_1(sK817) | m2_relset_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~v2_cat_1(sK817) | ~m2_cat_1(k11_isocat_2(sK817,sK818,sK819,sK820),sK817,sK818) <- (1515, 1517, 1521, 1532) [duplicate literal removal 151308]
151314. ~l1_cat_1(sK818) | ~v2_cat_1(sK817) | ~l1_cat_1(sK817) | ~m2_cat_1(k11_isocat_2(sK817,sK818,sK819,sK820),sK817,sK818) | v1_funct_2(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) <- (1515, 1517, 1521, 1532) [subsumption resolution 151311,44563]
151315. m2_relset_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~l1_cat_1(sK817) | ~v2_cat_1(sK817) | ~l1_cat_1(sK818) | ~m2_cat_1(k11_isocat_2(sK817,sK818,sK819,sK820),sK817,sK818) <- (1515, 1517, 1521, 1532) [subsumption resolution 151313,44563]
151316. ~l1_cat_1(sK817) | ~m2_cat_1(k11_isocat_2(sK817,sK818,sK819,sK820),sK817,sK818) | ~v2_cat_1(sK817) | ~l1_cat_1(sK818) <- (1515, 1517, 1521, 1532, ~1543) [subsumption resolution 151314,140999]
151317. ~m2_cat_1(k11_isocat_2(sK817,sK818,sK819,sK820),sK817,sK818) | m2_relset_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~v2_cat_1(sK817) | ~l1_cat_1(sK818) <- (1515, 1517, 1521, 1532) [subsumption resolution 151315,44562]
151318. ~m2_cat_1(k11_isocat_2(sK817,sK818,sK819,sK820),sK817,sK818) | ~l1_cat_1(sK818) | ~v2_cat_1(sK817) <- (1515, 1517, 1521, 1532, ~1543) [subsumption resolution 151316,44562]
151319. ~v2_cat_1(sK817) | ~l1_cat_1(sK818) | m2_relset_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) <- (1515, 1517, 1521, 1532) [subsumption resolution 151317,140573]
151320. ~v2_cat_1(sK817) | ~l1_cat_1(sK818) <- (1515, 1517, 1521, 1532, ~1543) [subsumption resolution 151318,140573]
151321. m2_relset_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) | ~l1_cat_1(sK818) <- (1515, 1517, 1521, 1532) [subsumption resolution 151319,44569]
151322. ~l1_cat_1(sK818) <- (1515, 1517, 1521, 1532, ~1543) [subsumption resolution 151320,44569]
151323. m2_relset_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) <- (1515, 1517, 1521, 1532) [subsumption resolution 151321,44564]
151324. $false <- (1515, 1517, 1521, 1532, ~1543) [subsumption resolution 151322,44564]
151325. ~1515 | ~1517 | ~1521 | ~1532 | 1543 [avatar contradiction clause 151324]
151326. 1547 | ~1515 | ~1517 | ~1521 | ~1532 [avatar split clause 151323,140921,140572,140451,140443,141014]
151374. m1_relset_1(k13_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK818)) <- (1547) [resolution 141015,65267]
151375. $false <- (~1544, 1547) [subsumption resolution 151374,141003]
151376. 1544 | ~1547 [avatar contradiction clause 151375]
151377. ~r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK819),u1_cat_1(sK817),u2_cat_1(sK819),k7_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) <- (~369, 1515, 1517) [forward demodulation 78354,140528]
151401. ~r4_nattra_1(u1_cat_1(sK817),u2_cat_1(sK819),u1_cat_1(sK817),u2_cat_1(sK819),k4_nattra_1(sK817,sK819,k12_isocat_2(sK817,sK818,sK819,sK820)),k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820))) <- (~369, 1515, 1517) [forward demodulation 151377,140596]
151407. ~1623 | 369 | ~1515 | ~1517 [avatar split clause 151401,140451,140443,78352,142420]
152048. ~v2_cat_1(sK817) | ~v2_cat_1(sK819) | ~l1_cat_1(sK817) | ~l1_cat_1(sK818) | ~l1_cat_1(sK819) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~v2_cat_1(sK818) | ~m2_nattra_1(k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),sK817,k11_cat_2(sK818,sK819),sK820,sK820) <- (~1620) [resolution 146696,142410]
152050. ~m2_nattra_1(k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),sK817,k11_cat_2(sK818,sK819),sK820,sK820) | ~l1_cat_1(sK817) | ~v2_cat_1(sK817) | ~l1_cat_1(sK819) | ~v2_cat_1(sK819) | ~l1_cat_1(sK818) | ~v2_cat_1(sK818) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) <- (~1620) [duplicate literal removal 152048]
152052. ~l1_cat_1(sK818) | ~v2_cat_1(sK819) | ~l1_cat_1(sK819) | ~l1_cat_1(sK817) | ~v2_cat_1(sK818) | ~v2_cat_1(sK817) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) <- (1515, 1517, ~1620) [subsumption resolution 152050,140535]
152054. ~l1_cat_1(sK819) | ~l1_cat_1(sK817) | ~v2_cat_1(sK817) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~v2_cat_1(sK819) | ~v2_cat_1(sK818) <- (1515, 1517, ~1620) [subsumption resolution 152052,44564]
152056. ~l1_cat_1(sK817) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~v2_cat_1(sK819) | ~v2_cat_1(sK818) | ~v2_cat_1(sK817) <- (1515, 1517, ~1620) [subsumption resolution 152054,44565]
152058. ~v2_cat_1(sK818) | ~v2_cat_1(sK817) | ~v2_cat_1(sK819) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) <- (1515, 1517, ~1620) [subsumption resolution 152056,44562]
152059. ~v2_cat_1(sK819) | ~v2_cat_1(sK817) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) <- (1515, 1517, ~1620) [subsumption resolution 152058,44563]
152060. ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~v2_cat_1(sK817) <- (1515, 1517, ~1620) [subsumption resolution 152059,44568]
152061. ~v2_cat_1(sK817) <- (1515, 1517, ~1620) [subsumption resolution 152060,44566]
152062. $false <- (1515, 1517, ~1620) [subsumption resolution 152061,44569]
152063. ~1515 | ~1517 | 1620 [avatar contradiction clause 152062]
152066. ~m2_nattra_1(k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),sK817,k11_cat_2(sK818,sK819),sK820,sK820) | ~l1_cat_1(sK817) | ~v2_cat_1(sK819) | ~v2_cat_1(sK817) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~v2_cat_1(sK818) | ~l1_cat_1(sK818) | ~l1_cat_1(sK819) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) <- (~1625) [resolution 142431,144155]
152067. ~v2_cat_1(sK818) | ~l1_cat_1(sK817) | ~v2_cat_1(sK817) | ~v2_cat_1(sK819) | ~l1_cat_1(sK818) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~l1_cat_1(sK819) | ~m2_nattra_1(k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),sK817,k11_cat_2(sK818,sK819),sK820,sK820) <- (~1625) [duplicate literal removal 152066]
152068. ~v2_cat_1(sK819) | ~l1_cat_1(sK818) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~l1_cat_1(sK817) | ~m2_nattra_1(k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),sK817,k11_cat_2(sK818,sK819),sK820,sK820) | ~v2_cat_1(sK817) | ~l1_cat_1(sK819) <- (~1625) [subsumption resolution 152067,44563]
152069. ~v2_cat_1(sK817) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~m2_nattra_1(k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),sK817,k11_cat_2(sK818,sK819),sK820,sK820) | ~l1_cat_1(sK817) | ~l1_cat_1(sK819) | ~l1_cat_1(sK818) <- (~1625) [subsumption resolution 152068,44568]
152070. ~l1_cat_1(sK818) | ~l1_cat_1(sK817) | ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~m2_nattra_1(k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),sK817,k11_cat_2(sK818,sK819),sK820,sK820) | ~l1_cat_1(sK819) <- (~1625) [subsumption resolution 152069,44569]
152071. ~m2_cat_1(sK820,sK817,k11_cat_2(sK818,sK819)) | ~m2_nattra_1(k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),sK817,k11_cat_2(sK818,sK819),sK820,sK820) | ~l1_cat_1(sK819) | ~l1_cat_1(sK817) <- (~1625) [subsumption resolution 152070,44564]
152072. ~l1_cat_1(sK819) | ~m2_nattra_1(k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),sK817,k11_cat_2(sK818,sK819),sK820,sK820) | ~l1_cat_1(sK817) <- (~1625) [subsumption resolution 152071,44566]
152073. ~l1_cat_1(sK817) | ~m2_nattra_1(k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),sK817,k11_cat_2(sK818,sK819),sK820,sK820) <- (~1625) [subsumption resolution 152072,44565]
152074. ~m2_nattra_1(k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820),sK817,k11_cat_2(sK818,sK819),sK820,sK820) <- (~1625) [subsumption resolution 152073,44562]
152075. $false <- (1515, 1517, ~1625) [subsumption resolution 152074,140535]
152076. ~1515 | ~1517 | 1625 [avatar contradiction clause 152075]
152102. m1_relset_1(k14_isocat_2(sK817,sK818,sK819,sK820,sK820,k4_nattra_1(sK817,k11_cat_2(sK818,sK819),sK820)),u1_cat_1(sK817),u2_cat_1(sK819)) <- (1625) [resolution 142430,65267]
152103. $false <- (~1622, 1625) [subsumption resolution 152102,142418]
152104. 1622 | ~1625 [avatar contradiction clause 152103]
152105. $false [avatar sat refutation 78359,140479,140493,140619,140674,140928,140937,140943,140975,141365,141440,142346,142423,144214,144215,144222,144236,146796,146798,146855,146873,151325,151326,151376,151407,152063,152076,152104]
% SZS output end Proof for CAT027+3
% ------------------------------
% Version: Vampire 4.9 (commit c3ff1b49a on 2025-01-24 06:41:25 +0100)
% Termination reason: Refutation

% >>> Input
% Input formulas: 11799
% >>> Preprocessing
% Introduced names: 527
% Introduced skolems: 3554
% Pure predicates: 25
% Unused predicate definitions: 2
% Function definitions: 61
% 
% >>> Saturation
% Initial clauses: 23188
% Generated clauses: 112316
% Activations started: 1011
% Active clauses: 1011
% Passive clauses: 48831
% Blocked clauses: 4
% 
% >>> Simplifying Inferences
% Duplicate literals: 1919
% Trivial inequalities: 82
% Fw subsumption resolutions: 10222
% Fw demodulations: 24148
% Bw demodulations: 27054
% 
% >>> Deletion Inferences
% Simple tautologies: 225
% Equational tautologies: 46
% Forward subsumptions: 25463
% Fw demodulations to eq. taut.: 9
% Bw demodulations to eq. taut.: 134
% 
% >>> Generating Inferences
% Binary resolution: 2607
% Forward superposition: 420
% Backward superposition: 816
% Self superposition: 1
% 
% >>> AVATAR
% Split clauses: 1711
% Split components: 4150
% Sat splitting refutations: 399
% 
% >>> SAT Solver Statistics
% SAT solver clauses: 2520
% SAT solver unit clauses: 141
% SAT solver binary clauses: 815
% 
% >>> Neural Instruction Counts
% Neural model warmup: 473
% Gnn eval: 28096
% Bulk evals: 9493
% 
% Time elapsed: 21.057 s
% Peak memory usage: 830 MB
% Instructions burned: 52658 (million)
% ------------------------------
% ------------------------------
